Mid-zone hepatocytes trade proliferation for survival via Atf4-Chop axis in early acute liver injury

Author response:

The following is the authors’ response to the previous reviews.

eLife Assessment

This study addresses an important question in liver biology: how zonal hepatocytes balance survival and proliferation following injury? The authors propose that a midzone Atf4-Chop axis to Btg2 program temporarily suppresses proliferation to promote survival after a variety of chemical and surgical liver injury models. The authors provide evidence that some zones mount tailored stress responses, which ultimately promote regeneration and liver healing; however, the "mid-zone" changes with different injury models, making it difficult to conclude that the ATF4-CHOP response is specific to this zone in all injury contexts. In addition, it is possible that Atf4 and Btg2 overexpression could lead to Cyp2e1 suppression, which could reduce the extent of injury after CCl4 or APAP. To some extent, these points make the strength of the evidence incomplete, but do not entirely detract from the significance of the study, which is underscored by the helpful observation that there are zone-specific stress responses that mediate liver regeneration and survival.

We appreciate the editor’s thoughtful assessment. Regarding regional specificity, we agree that "mid-zone" patterns differ across injury models, particularly in HPx; to maintain consistency, we have removed the PHx data and now focus our conclusions on the APAP and CCl4 models. Concerning the potential confound of Cyp2e1 suppression, our AAV overexpression resulted in only modest reductions in Cyp2e1 protein (Revised Figure 4-figure supplement 3A). Although enzymatic activity was not directly measured, basal CYP2E1 protein levels generally correlate with enzymatic activity. Furthermore, published studies indicate that protection from APAP toxicity requires >50% CYP2E1 suppression (PMID: 35145060; PMID: 30151903), whereas the reduction observed here was substantially smaller and therefore unlikely to account for the marked decrease in liver injury.

Reviewer #2 (Public review):

The manuscript reports protection of midlobular hepatocytes from APAP toxicity by activation of Atf4-CHOP (Ddit3)-mediated cell cycle arrest and stress response. The authors acknowledge that their finding is unexpected because CHOP typically induces cell death. Therefore, they functionally validate several aspects of the proposed Atf4-CHOP mechanism. Along these lines, the mitigation of APAP toxicity by AAV expression of Atf4 or Btg2, the latter identified as CHOP effector, is impressive. Whether Atf4 indeed acts through CHOP and whether midlobular hepatocytes are protected because of cell cycle arrest is less clear. These and other criticisms are described in the following.

Major points:

(1) The difference in baseline Cyp2e1 expression between F2A and F3B remains unexplained after revision.

We apologize for the confusion. The apparent discrepancy reflects different visualization methods rather than a biological inconsistency. Figure 2A and Figure 3B are derived from the same spatial transcriptomic dataset but are visualized using different normalization strategies. Figure 2A uses row-wise Z-score normalization to emphasize the relative spatial distribution of each gene across liver zones. Consequently, even a moderate spatial gradient (e.g., ~1.5–2-fold) is visually enhanced to facilitate comparison of zonation patterns between genes. In contrast, Figure 3B presents the absolute log-normalized transcript abundance without rowwise scaling. Because Cyp2e1 is highly expressed in both the pericentral and adjacent mid-zone hepatocytes, the absolute difference between these regions is relatively modest, resulting in a less pronounced visual contrast. Thus, the different appearance of Cyp2e1 between Figures 2A and 3B reflects the visualization strategy rather than conflicting expression data. To avoid further confusion, we have added this clarification to the legends of both Figures 2A and 3B.

(2) In contrast to the revised discussion, the abstract does not reflect that limited evidence for a cell cycle arrest in pericentral hepatocytes was found.

We thank the reviewer for the continued effort in improving our manuscript. To reflect the limited support of ST data for cell cycle arrest, we have revised the Abstract (Revised manuscript, page 1, line 13-26)

(3) Additional Btg2 knockout data support its proposed role in the revised manuscript.

We appreciate the reviewer’s acknowledgement of the new functional data supporting the role of Btg2.

(4) The BTG2 immunostaining remains weak, not only in in F6F but now also in F6D of the revised manuscript, which together with lack of high-resolution immunostaining of AAV-Ddit3-induced BTG2 in the absence of APAP results in limited support for the conclusion that APAP promotes nuclear localization of BTG2.

We agree that the current immunostaining does not provide definitive evidence for nuclear localization. In response, we have deliberately avoided any claims regarding nuclear translocation in the revised manuscript, instead framing our findings strictly around increased BTG2 expression and accumulation in the peri-necrotic mid-zone region following APAP injury.

(5) The extended list of transcription factors (from 30 to 50) includes Atf4 but direct evidence for an interaction with Ddit3 is missing from the revised manuscript.

We thank the reviewer for this important point. We agree that our current dataset does not provide direct evidence for a physical interaction between Atf4 and Ddit3 proteins. Throughout the manuscript, we use the term "Atf4–Chop axis" to denote a functional regulatory pathway rather than a direct protein–protein interaction. To avoid ambiguity, we have revised the manuscript to characterize Atf4 and Ddit3 as co-activated transcriptional co-regulators and to restrict our conclusions to transcriptional co-regulation and pathway convergence, as supported by SCENIC, Cut&Run, and GO analyses. We now explicitly define "axis" as a functional pathway (revised manuscript, page 5, line 182–184; page 6, line 260–261; page 8, line 331333).

(6) The ATF4 immunostaining after APAP challenge remains weak.

We agree that the endogenous ATF4 immunostaining is relatively weak. This is consistent with the low basal abundance and transient induction of ATF4 protein during the integrated stress response. To improve visualization, we have added higher-magnification insets to Revised Figure 5A, which more clearly show nuclear ATF4 staining in hepatocytes adjacent to the necrotic region following APAP treatment. We hope these enlarged images better illustrate the spatial distribution of ATF4 while accurately reflecting its endogenous expression level.

(7) S5A of the revised manuscript rules out loss of Cyp2e1 expression as a confounding factor.

We thank the reviewer for acknowledging that Revised Figure S5A addresses this concern.

(8) The revised manuscript continues to focus on rare spatial transcriptomics analyses of patients with APAP toxicity although more snRNA-seq analyses of such patients are available which should also allow for analysis of hepatocyte zonation.

We thank the reviewer for this constructive suggestion and apologize for not describing our analyses more clearly in the previous revision. In response, we analyzed both the snRNA-seq and spatial transcriptomics (ST) datasets from GSE223561, which includes snRNA-seq profiles from 9 healthy and 10 APAP explants together with ST data from 3 healthy and 2 APAP patients (Revised Figure 4-figure supplement 3).

The snRNA-seq data readily resolved hepatocyte zonation using established pericentral (Glul, Cyp2e1, Cyp2a5), periportal (Alb, Cyp2f2, Sds), and mid-zonal (Igfbp2, Hamp) markers (Figure 1A–D). Consistent with our mouse data, mid-zonal hepatocytes in APAP patients showed increased expression of stress-response genes (DDIT3, ATF4, and HMOX1). However, unlike the early mouse injury model, proliferation-associated genes (MKI67 and CCNB1) were also elevated (Revised Figure 4-figure supplement 3E), likely reflecting the end-stage nature of liver explants, in which regenerative responses are already well established rather than the early (3–12 h) injury phase examined in our study.

The ST dataset provided complementary spatial information that cannot be obtained from dissociated nuclei alone. Although limited in sample number, one APAP specimen closely recapitulated our murine observations, showing mid-zonal stress enrichment accompanied by reduced proliferation, whereas the second exhibited a predominantly proliferative profile consistent with a later regenerative stage (Revised Figure 4-figure supplement 3F).

Together, these analyses indicate that human mid-zonal hepatocytes exhibit a conserved stress-response program, while the relationship between stress signaling and proliferation is highly dependent on the stage of injury. We have incorporated both the snRNA-seq and ST analyses into the revised manuscript and explicitly discuss the temporal limitations of the available human datasets (page 5-6, lines 225239; page 9, lines 368-379).

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

(1) F2A includes Ddit3 in the revised manuscript and although it is not significant in S1G (former F2B), it is among the most highly expressed transcription factors in F4B and S3B.

We thank the reviewer for this careful observation. As noted, Ddit3 is differentially expressed in the mid-zonal hepatocytes at both 3 and 6 h after APAP treatment (Revised Figure 2A). We also agree that Ddit3 does not reach statistical significance in the differential expression analysis shown in Revised Figure 2-figure supplement 1G under our predefined thresholds. However, our rationale for prioritizing Ddit3 was based primarily on transcription factor activity rather than differential expression alone. Differential expression analysis evaluates changes in the abundance of an individual transcript, whereas SCENIC regulon analysis infers the functional activity of a transcription factor from the coordinated expression of its downstream target genes. Consequently, transcription factors with relatively modest changes in mRNA abundance may nevertheless exhibit high regulon activity if their downstream regulatory network is broadly activated.

Consistent with this, Ddit3 was identified as the highest-activity regulon in the mid zone, while Atf4 ranked seventh (Figures 4B and Figure 4-figure supplement 1B). Given the well-established cooperative roles of ATF4 and DDIT3 in the integrated stress response, these unbiased network analyses identified the ATF4–DDIT3 axis as a biologically relevant pathway for further investigation. To avoid overinterpretation, we have revised the manuscript to clarify that DDIT3 was prioritized based on its consistently high regulon activity and the activation of its downstream transcriptional network, rather than on differential expression alone (Revised manuscript, page 5, line 188-192).

We thank the reviewers for their rigorous critique again. We thank eLife for fostering an environment of fairness and transparency that enables authors to communicate openly and present their data honestly.

Changes from the earlier version to the revised version

Abstract

Earlier version: Hepatocytes undergo extensive proliferation to facilitate liver repair after injury, yet early adaptive changes prior to proliferation remain unclear. Here, we report that during early acetaminophen (APAP)-induced liver injury, hepatocytes exhibit transient proliferation suppression, most pronounced in mid-zone hepatocytes due to zonal APAP metabolism. Using spatial transcriptomics (ST), immunohistochemistry, and functional studies, we identified a unique mid-zone stress-response program. Central to this adaptation is the Atf4-Chop axis, which actively suppresses proliferation via the cell cycle inhibitor Btg2, prioritizing cytoprotection over cell division. This transient arrest is a critical survival strategy: halting energy-intensive proliferation during peak injury allows mid-zone hepatocytes to redirect resources towards protection, enhancing their survival in early APAP-induced liver injury. Thus, Atf4-Chop-mediated quiescence preserves a hepatocyte reservoir necessary for subsequent regenerative proliferation and effective repair. Our findings reveal a key adaptive trade-off in mid-zone hepatocytes where transient proliferation arrest promotes early survival to enable repair.

Revised version: Hepatocytes undergo extensive proliferation to facilitate liver repair after injury, yet early adaptive changes prior to proliferation remain unclear. Here, we report that during early acetaminophen (APAP)-induced liver injury, hepatocytes exhibit transient proliferation suppression that is most pronounced in mid-zone hepatocytes, consistent with zonal APAP metabolism. While spatial transcriptomics (ST) provided robust evidence for this arrest in the mid-zone, support for a similar arrest in pericentral hepatocytes was limited. Integrating ST with immunohistochemistry and functional studies, we identified a unique mid-zone stress-response program centered on the Atf4–Chop axis, which suppresses proliferation via the cell cycle inhibitor Btg2. Together, our findings support a model in which mid-zone hepatocytes transiently prioritize stress adaptation over proliferation, thereby preserving regenerative capacity for subsequent liver repair.

Introduction

Earlier version: Acetaminophen (APAP) overdose is the leading cause of drug-induced acute liver failure worldwide, resulting in over 10,000 hospitalizations and approximately 500 deaths annually in the United States (Bernal and Wendon, 2013; Larson et al., 2005; Lee, 2007; Ostapowicz et al., 2002). The hepatotoxicity of APAP is initiated by its metabolic activation through cytochrome P450 enzymes, primarily in pericentral hepatocytes, generating the reactive metabolite N-acetyl-p-benzoquinone imine (NAPQI) (Dahlin et al., 1984; Hinson et al., 2004; Hinson et al., 2010; Tujios and Fontana, 2011). NAPQI depletes glutathione and forms protein adducts, leading to oxidative stress and mitochondrial dysfunction that culminate in centrilobular necrosis (Dahlin et al., 1984; Hinson et al., 2004; Hinson et al., 2010; Tujios and Fontana, 2011). Although the mechanisms of APAP bioactivation and toxicity are well-characterized, the early adaptive responses of hepatocytes—particularly their proliferative dynamics across lobular zones—remain poorly understood.

Revised version: Acetaminophen (APAP) overdose is the leading cause of drug-induced acute liver failure worldwide, resulting in over 10,000 hospitalizations and approximately 500 deaths annually in the United States (Ostapowicz, Fontana et al. 2002, Larson, Polson et al. 2005, Lee 2007, Bernal and Wendon 2013). The hepatotoxicity of APAP is initiated by its metabolic activation through cytochrome P450 enzymes, primarily in pericentral hepatocytes, generating the reactive metabolite N-acetyl-p-benzoquinone imine (NAPQI) (Dahlin, Miwa et al. 1984, Hinson, Reid et al. 2004, Jack A. Hinson 2010, Tujios and Fontana 2011). NAPQI depletes glutathione and forms protein adducts, leading to oxidative stress and mitochondrial dysfunction that culminate in centrilobular necrosis (Dahlin, Miwa et al. 1984, Hinson, Reid et al. 2004, Jack A. Hinson 2010, Tujios and Fontana 2011). Although the mechanisms of APAP bioactivation and toxicity are well-characterized, the early adaptive responses of hepatocytes—particularly their proliferative dynamics across lobular zones—remain poorly understood.

Earlier version: The liver is organized into functional hexagonal units termed lobules, where hepatocytes display spatial heterogeneity across three distinct zones: periportal (PP, zone 1), mid-zonal (Mid, zone 2), and pericentral (PC, zone 3) (Ben-Moshe et al., 2022a; Wang et al., 2024; Wu et al., 2024). These zones exhibit unique gene expression profiles and metabolic specializations(Halpern et al., 2017). Pericentral hepatocytes (zone 3), expressing high levels of cytochrome P450 enzymes, are the primary site of APAP-induced necrosis (Wang et al., 2024). However, recent spatial profiling, functional, and lineage tracing studies have identified midzonal hepatocytes (zone 2) as the predominant source of new hepatocytes during both homeostasis and repair, demonstrating significant proliferative capacity and metabolic plasticity (Chembazhi et al., 2021; He et al., 2021; Wei et al., 2021). Despite these advances, how mid-zone hepatocytes dynamically adapt during the initial phase of acute stress remains unknown.

Revised version: The liver is organized into functional hexagonal units termed lobules, where hepatocytes display spatial heterogeneity across three distinct zones: periportal (PP, zone 1), mid-zonal (Mid, zone 2), and pericentral (PC, zone 3) (Ben-Moshe, Veg et al. 2022, Wang, Wang et al. 2024, Wu, Shentu et al. 2024). These zones exhibit unique gene expression profiles and metabolic specializations(Halpern, Shenhav et al. 2017). Pericentral hepatocytes (zone 3), expressing high levels of cytochrome P450 enzymes, are the primary site of APAP-induced necrosis (Wang, Wang et al. 2024). However, recent spatial profiling, functional, and lineage tracing studies have identified midzonal hepatocytes (zone 2) as the predominant source of new hepatocytes during both homeostasis and repair, demonstrating significant proliferative capacity and metabolic plasticity (Chembazhi, Bangru et al. 2021, He, Pu et al. 2021, Wei, Wang et al. 2021). Despite these advances, how mid-zone hepatocytes dynamically adapt during the initial phase of acute stress remains unknown.

Earlier version: A key regulator of cellular stress adaptation is the integrated stress response (ISR), which phosphorylates eukaryotic initiation factor 2α (eIF2α) to suppress global translation while selectively upregulating stress-responsive genes, including activating transcription factor 4 (Atf4) and its downstream target pro-apoptotic factor C/EBP homologous protein (Chop, encoded by DNA damage-inducible transcript 3 [Ddit3]) (Costa-Mattioli and Walter, 2020; Vasudevan et al., 2020). Although Chop is traditionally associated with apoptosis, its transcriptional targets exhibit significant overlap with those of Atf4, encompassing genes that paradoxically support cell survival and growth (Han et al., 2013; Marciniak et al., 2004; Uzi et al., 2013). These findings underscore the context-dependent duality of the Atf4-Chop axis across different stress conditions, highlighting the need to elucidate its precise regulatory roles in specific physiological settings.

Revised version: A key regulator of cellular stress adaptation is the integrated stress response (ISR), which phosphorylates eukaryotic initiation factor 2α (eIF2α) to suppress global translation while selectively upregulating stress-responsive genes, including activating transcription factor 4 (Atf4) and its downstream target pro-apoptotic factor C/EBP homologous protein (Chop, encoded by DNA damage-inducible transcript 3 [Ddit3]) (Costa-Mattioli and Walter 2020, Vasudevan, Neuman et al. 2020). Although Chop is traditionally associated with apoptosis, its transcriptional targets exhibit significant overlap with those of Atf4, encompassing genes that paradoxically support cell survival and growth (Marciniak, Yun et al. 2004, Han, Backa et al. 2013, Uzi, Barda et al. 2013). These findings underscore the context-dependent duality of the Atf4-Chop axis across different stress conditions, highlighting the need to elucidate its precise regulatory roles in specific physiological settings.

Through spatial transcriptomics (ST) and functional analyses, we demonstrate that zonal heterogeneity in APAP metabolism drives a stress-adaptive response in mid-zone hepatocytes during early APAP-induced liver injury. This response is mediated by the Atf4-Chop-Btg2 axis, which orchestrates a transient proliferation arrest to prioritize cellular survival over regenerative capacity. Our findings elucidate the spatial and molecular mechanisms governing hepatocyte adaptation to acute injury, uncovering a critical survival-regeneration trade-off in early phase of acute liver injury.

Results

Mid-zone hepatocytes exhibit the most pronounced proliferative decline during early acute liver injury

Earlier version: To assess hepatocyte proliferation in early APAP-induced liver injury, wild-type C57BL/6J mice were intraperitoneally injected with 300 mg/kg APAP, and liver sections were analyzed at multiple time points (0, 3, 6, 12, and 24 hours post-APAP). Immunohistochemical staining for Ki67, a proliferation marker, revealed that mid-zone hepatocytes exhibited a roughly 3-fold higher basal proliferative activity than PP and PC hepatocytes under homeostatic conditions (0 hour) (Figure 1A, B), which in line with previous study (Wei et al., 2021). During the early injury phase (3–12 hours post-APAP), the number of Ki67+ hepatocytes in the mid-zone sharply declined from 70 to 20 per section, reaching proliferation levels similar to those in PP and PC zones (Figure 1A, B). By 24 hours post-APAP, mid-zone proliferation gradually recovered (Figure 1A, B), indicating a transient suppression of proliferation in this region during early APAP-induced liver injury. Further analysis of proliferation dynamics in each zone confirmed that all zones exhibited reduced proliferation during the injury initiation phase, with the mid-zone showing the most decline (Figure 1C). Comparative value change analysis (relative to 0 hour) further supported these findings, highlighting the mid-zone as the most affected region in terms of proliferative suppression during early APAP-induced liver injury (Figure 1D).

Revised version: To assess hepatocyte proliferation during early APAP-induced liver injury, wild-type C57BL/6J mice received an intraperitoneal injection of 300 mg/kg APAP, and liver sections were analyzed at 0, 3, 6, 12, and 24 hours post-injection. Immunohistochemical staining for the proliferation marker Ki67 revealed that under homeostatic conditions (0 h), mid-zone hepatocytes exhibited approximately three-fold higher basal proliferative activity than periportal (PP) and pericentral (PC) hepatocytes, consistent with previous reports (Wei, Wang et al. 2021). During the early injury phase (3–12 h), the number of Ki67⁺mid-zone hepatocytes declined sharply from roughly 70 to 20 per section, reaching levels comparable to those observed in the PP and PC zones (Figure 1A, B). By 24 h post-APAP, mid-zone proliferation gradually recovered toward baseline levels (Figure 1A, B). Although all zones exhibited reduced proliferation during the injury initiation phase, the mid-zone showed the most substantial decline (Figure 1C), as further corroborated by value-change analysis relative to the 0 h time point (Figure 1D). Collectively, these data identify the mid-zone as the region most profoundly affected by transient proliferative suppression during early APAP-induced acute liver injury.

Mid-zone hepatocytes exhibit the most pronounced proliferative decline in early acute liver injury.

Wildtype C57BL/6J mice were intraperitoneally injected with 300 mg/kg APAP for 0, 3, 6, 12 and 24 h. (A) Immunohistochemical staining of Ki67 to detect proliferating hepatocytes in liver sections at 0, 3, 6, 12 and 24 h post-APAP. Representative images of pericentral (PC), middle (Mid) and periportal (PP) are shown. Red arrows indicate Ki67+ hepatocytes. (B) Quantitative analysis of Ki67+ hepatocytes across liver zones (PC, Mid, PP) at 0, 3, 6, 12, and 24 h post-APAP. (C) Quantitative analysis of Ki67+ hepatocytes across 0, 3, 6, 12, and 24 h post-APAP in each liver zone (PC, Mid, PP). (D) Changes of Ki67+ hepatocytes across liver zones (PC, Mid, PP) at 0, 3, 6, 12 and 24 h post-APAP. Value change=number of Ki67+ hepatocytes at x h post-APAP - the number of Ki67+ hepatocytes at 0 h post-APAP. (a) denotes significance between PC and Mid regions, (b) denotes significance between PC and PP regions, and (c) denotes significance between Mid and PP regions. Data points represent mean hepatocyte counts from six high-power fields per zonal layer (n=3 mice/group) in B and C. Data are represented as means ± SD; One-way ANOVA (B,C,D).*p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Mid-zone hepatocytes show distinct gene expression profiles

Earlier version: To investigate the mechanism underlying proliferation pause during early APAP-induced liver injury, liver sections from wild-type mice treated with APAP for various time points were subjected to ST using the Visium (10 × Genomics) platform (Figure S1A). Due to variability in mice’s responses to APAP, multiple mice were treated at each time point and one sample representing the average level at each time point was selected based on Hematoxylin and eosin (H&E) staining and serum alanine aminotransferase (ALT) levels for ST analysis. H&E staining of liver sections used for ST revealed a gradual increase in the necrotic area from 0 to 24 hours (Figure S1B). Under normal conditions (APAP 0h), each liver lobule was divided into three zones based on unbiased gene expression profiles. The PP zone was defined by enrichment of PP signature genes (e.g., Alb, Mup20, Cyp2f2, Pck1, Apoa4). The_PC zone was defined by high expression of PC markers (e.g., Glul, Cyp2e1, Oat, Cyp1a2, Apoe). The Mid zone comprised regions with intermediate expression of PC and PP markers and elevated levels of Igfbp2 and Hamp (Figure 2A and S1C). Following APAP-induced injury (3h, 6h), the PC zone remained identifiable based on residual enrichment of PC signature genes (e.g., Cyp1a2, Glul) despite necrosis and reduced overall transcription, while PP gene expression remained largely unchanged. The Mid zone was defined as the transcriptional cluster between PC and PP regions exhibiting marked reprogramming, (e.g., Sqstm1, Igfbp1) (Figure 2A and S1C). Analysis of ST data at 0, 3, 6, and 24 hours post-APAP via UMAP plots consistently revealed distinct zonal clusters and dynamic spatiotemporal patterns during the early phase of acute injury (Figures S1D and S1E). Spatiotemporally resolved heatmaps of representative markers, such as Glul (a PC marker) and Cyp2f2 (a PP marker), further reinforced the stability of zonation profiles during early acute liver injury (Figure S1F).

Revised version: To investigate the mechanism underlying proliferation pause during early APAP-induced liver injury, liver sections from wild-type mice treated with APAP for various time points were subjected to ST using the Visium (10 × Genomics) platform (Figure 2-figure supplement 1A). Given inter-mouse variability, one representative sample per time point was selected based on H&E histology and serum ALT levels (Figure 2-figure supplement 1B). Under normal conditions (APAP 0h), unbiased transcriptional clustering partitioned each lobule into three distinct zones: PP (Alb, Mup20, Cyp2f2), PC (Glul, Cyp2e1, Cyp1a2), and mid-zone (intermediate between PP and PC, with elevated Igfbp2 and Hamp) (Figure 2A, Figure 2-figure supplement 1C). Following APAP-induced injury (3h, 6h), the PC zone remained identifiable by residual PC signature genes despite necrosis and reduced transcription, while PP genes remained largely unchanged. The Mid zone was distinguishable as the intervening cluster exhibiting marked transcriptional reprogramming (e.g., Sqstm1, Igfbp1) (Figure 2A and Figure 2-figure supplement 1C). UMAP projections and spatial heatmaps of Glul and Cyp2f2 confirmed persistent, though dynamic, zonal clustering across all time points (Figure 2-figure supplement 1D--F). These data indicate that the fundamental lobular architecture remains spatially resolved during the early injury phase, enabling zone-specific transcriptional analysis.

Mid-zone hepatocytes show distinct gene expression profiles.

Earlier version: (A) Heatmap of the top 10 differentially expressed genes in periportal (PP), midzonal (Mid) and pericentral (PC) hepatocyte zones at 0, 3, 6 and 24 h post-APAP. Functional categories at the bottom were assigned using Gene Ontology (GO) terms (biological process, cellular component, molecular function); for genes annotated with multiple GO terms, the most representative or significantly enriched term was selected. (B) Immunofluorescence staining for glutamine synthetase (GS, a pericentral marker) with DAPI (blue) in liver sections at 0, 3, 6, and 24 h after APAP administration. Quantification of the zonal distribution of GS-positive cells at the indicated time points. (C) Immunohistochemical detection of p62 (Sqstm1 protein, a midzone marker during early injury) in liver sections at 0, 3, 6, and 24 h after APAP administration. Red arrows indicate representative p62-positive hepatocytes. Quantification of the zonal distribution of p62-positive cells at the indicated time points. (D) Immunohistochemical staining of Cyp2f2 (a periportal marker) in liver sections at 0, 3, 6, and 24 h after APAP administration. Quantification of the zonal distribution of Cyp2f2-positive cells at the indicated time points. Data are presented as mean ± SD. One-way ANOVA (B-D). *p < 0.05, **p < 0.01, ***p < 0.001, ****p < 0.0001; ns, not significant.

Revised version: (A) Heatmap of the top 10 differentially expressed genes in periportal (PP), midzonal (Mid) and pericentral (PC) hepatocyte zones at 0, 3, 6 and 24 h post-APAP. To highlight relative expression shifts across zones, this heatmap is visualized using row-wise Z-score standardization, which scales each gene to a uniform range and emphasizes spatial switching patterns rather than absolute transcript abundance. Functional categories at the bottom were assigned using Gene Ontology (GO) terms (biological process, cellular component, molecular function); for genes annotated with multiple GO terms, the most representative or significantly enriched term was selected. (B) Immunofluorescence staining for glutamine synthetase (GS, a pericentral marker) with DAPI (blue) in liver sections at 0, 3, 6, and 24 h after APAP administration. Quantification of the zonal distribution of GS-positive cells at the indicated time points. (C) Immunohistochemical detection of p62 (Sqstm1 protein, a midzone marker during early injury) in liver sections at 0, 3, 6, and 24 h after APAP administration. Red arrows indicate representative p62-positive hepatocytes. Quantification of the zonal distribution of p62-positive cells at the indicated time points. (D) Immunohistochemical staining of Cyp2f2 (a periportal marker) in liver sections at 0, 3, 6, and 24 h after APAP administration. Quantification of the zonal distribution of Cyp2f2-positive cells at the indicated time points. Data are presented as mean ± SD. One-way ANOVA (B-D). *p < 0.05, **p < 0.01, ***p < 0.001, ****p < 0.0001; ns, not significant.

Added in the revised version: Differential expression analysis across zones and time points revealed distinct zonal trajectories (Figure 2A). In the PP zone, gene expression remained relatively stable across 0–6 h, with inflammatory genes (Orm1, Serpina1c) rising only at 24 h. The mid-zone exhibited a clear temporal progression. At 0 h, it expressed proliferative (Igfbp2), metabolic (Scd1, Cdo1), and homeostatic (Hamp, Cps1) genes. By 3–6 h, these were replaced by stress-response and autophagy genes (Hmox1, Hspa8, Atf3, Ubc, Sqstm1). By 24 h, wound-healing genes (Fgg, Fgl1) became prominent. The PC zone showed early downregulation of xenobiotic metabolism (Cyp2e1, Cyp1a2), followed by upregulation of cytoskeletal (Actb, Krt8) and redox (Prdx1, Srxn1) genes at 24 h. This temporal trajectory identifies the mid-zone as the most transcriptionally dynamic region, undergoing a functional switch from homeostasis to active stress adaptation during the initial injury phase.

Added in the revised version: Focusing on the mid-zone at 6 h versus 0 h, volcano plot analysis identified significant upregulation of stress-related genes (Atf3, Hmox1, Ddit3, Atf4) alongside cell-cycle and apoptosis regulators (Btg2, Egr1, Sqstm1) (Figure 2-figure supplement 1G). Among these, we selected Sqstm1 (encoding p62) for validation as a representative stress-response gene with previously characterized roles in autophagy and oxidative stress. To validate our zonation assignment, we compared ST clusters with a nine-layered radial classification from the central vein (CV) to the portal vein (PV). Immunostaining for GS (PC), p62 (mid-zone), and Cyp2f2 (PP) confirmed that our PC, Mid, and PP zones corresponded to layers 1–2, 3–6, and 7–9, respectively (Figure 2B--D). These spatial correlations confirm that the mid-zone is the primary site of early stress-response gene induction, with p62 serving as a robust regional marker.

Added in the revised version: ST-based analysis of Ki67+ spots recapitulated the transient proliferative reduction at 3-6 h (Figure 2-figure supplement 1I). To further dissect cell-cycle dynamics, we assessed S-phase and G2/M-phase gene signatures. Mid-zone hepatocytes displayed a transient increase in S-phase scores at 3 h, followed by a decline at 6 and 24 h; conversely, G2/M scores were reduced at 3-6 h and returned to baseline by 24 h (Figure 2-figure supplement 1J). Collectively, these ST analyses establish that mid-zone hepatocytes adopt a spatially confined stress-adaptive transcriptional program—marked by p62 upregulation and a impaired cell-cycle progression—during the early phase of APAP-induced injury, providing a molecular framework for the proliferative pause observed histologically.

Removed in the revised version: We further analyzed differentially expressed genes (DEGs) in hepatocyte zones (PP, Mid, and PC) at 0, 3, 6, and 24 hours post-APAP (Figure 2A). Throughout the 0 to 24-hour period post-APAP, gene expression in the PP remained relatively stable. At 0, 3, and 6 hours post-APAP, genes associated with metabolism (such as Hsd17b13, Fabp1, Apoc2, Thrsp for fat metabolism; Rida, Hal, Gls2, Bhmt for fatty acid metabolism; Aldob, Pck1 for glucose metabolism) and transport (such as Alb, Hpx, Selenop, Trf, Lcn2) were prominently expressed, with a notable upregulation of inflammatory response-related genes (such as Orm1, Serpina1c, Lrg1) observed at 24 hours post-APAP (Figure 2A). In contrast, Mid-zone exhibited significant changes over time. Initially, it showed high expression of genes related to cell proliferation (such as Igfbp2), metabolism (such as Scd1, Cdo1, Ces1c, Ttr, Apoa1), and homeostasis (such as Hamp, Cps1, Rgn) at 0 hours, transitioning to stress response/cell survival (such as Hmox1, Hspa8, Atf3, Osgin1) and protein degradation/autophagy (such as Ubb, Ubc, Sqstm1) at 3 and 6 hours, and eventually shifting to cell proliferation and wound healing (such as Fgg, Fgl1, Fgb) by 24 hours post-APAP (Figure 2A). The PC showed a decrease in xenobiotic metabolism (such as Cyp3a11, Cyp2e1, Cyp1a2, Cyp2c37) and fat metabolism (such as Apoe, Akr1c6, Hpd, Oat) during early APAP-induced liver injury (at 0, 3, 6 hours post-APAP), with a marked increase in the expression of genes related to cytoskeleton (such as Actb, Krt8, Krt18) and redox regulation (such as Prdx1, Srxn1) observed at 24 hours post-APAP (Figure 2A). These differential expression patterns across various zones and time points post-APAP highlight significant changes in Mid-zone. Notably, our analysis suggests that the Mid-zone undergoes a stress adaption during the initial phases of acute liver injury.

Removed in the revised version: Further analysis focuses on DEGs in Mid-zone at 6 hours post-APAP compared to 0 hour. The volcano plot revealed a notable increase in the number of genes significantly upregulated in the mid zones at 6 hours, predominantly associated with the stress response and adaptation (such as Atf3, Hmox1, Osgin1, Ddit3, Gadd45a, Atf4, Jun, Hspa1) as well as cell proliferation and apoptosis (such as Btg2, Egr1, Igfbp1, Sqstm1, Dnajb1) (Figure S1G). Among the significantly upregulated genes, Sqstm1 (Sequestosome 1) was selected for further validation. Sqstm1 encodes the p62 protein, a key regulator of autophagy and oxidative stress response (Huang et al., 2023). While Sqstm1 expression was almost evenly distributed across all zones at 0 hours, it was markedly elevated in the mid-zone at 3 and 6 hours post-APAP (Figure S1H). To validate and quantify our zonation approach during early injury, we compared the assigned zones with a classical division of the lobule into nine even layers spanning from the central vein (CV) to the portal vein (PV). Immunostaining and quantification for glutamine synthetase (GS, the protein product of Glul, a PC marker) and immunohistochemical staining for p62 (a midzone marker during early injury) and Cyp2f2 (a PP marker) further corroborated the zone definitions at each time point, demonstrating that our PC, Mid, and PP zones corresponded to layers 1–2, 3–6, and 7–9, respectively (Figure 2B–D). These findings suggest that while gene expression profiles change during acute injury, the fundamental spatial architecture of liver zones remains stable.

Removed in the revised version: Further ST analysis revealed that Ki67-positive spots also show a transient reduction during the early injury (at 3 and 6 hours post-APAP) (Figure S1I), which aligns with the findings from Ki67 immunohistochemical staining. Cell cycle analysis is crucial for understanding cell proliferation, as it reveals the stages of the cell cycle that cells are in and how many cells are actively dividing (Woo et al., 2009). In the cell cycle, the S phase is responsible for DNA synthesis and replication, while the G2M phase prepares cells for and executes mitosis (Woo et al., 2009). Analysis of cell-cycle phase distribution, based on genes identified in previous studies (Bunch et al., 2023; Ohoka et al., 2005), revealed that Mid-zone hepatocytes exhibited a transient increase in S-phase genes at 3 hours, followed by a decline at later time points (6, and 24 hours). In contrast, Mid-zone hepatocytes exhibited reduced G2/M-phase scores at 3 and 6 hours, returning to baseline levels by 24 hours (Figure S1J). This suggests that Mid-zone hepatocytes may initiate DNA synthesis, but they are unable to complete the cycle due to the environmental toxicity or cellular damage. Nuclear Autoantigenic Sperm Protein (Nasp) binds to histones and facilitate chromatin assembly during the S phase (Richardson et al., 2006). Casein Kinase 1 Beta (Cks1b) participates in the control of G1 phase and the transition to the S phase (Peters et al., 2000). High Nasp levels indicate efforts toward DNA repair and chromatin stability, while low Cks1b levels point to a potential cell cycle arrest or a reduction in cell proliferation to manage the damage (Figure S1K). Overall, ST analysis revealed that Mid-zone hepatocytes exhibit distinct gene expression profiles during early acute injury.

Zonal metabolism of APAP leads to stress response in Mid-zone hepatocytes

Added in the revised version: To determine what triggers the mid-zone stress response, we first examined the spatial distribution of APAP-metabolizing cytochrome P450 (Cyp) enzymes. Cyp enzymes metabolize APAP to NAPQI, which causes oxidative stress and cellular damage unless neutralized by GSH (Figure 3A). ST analysis revealed a zonal gradient in the expression of most Cyp transcripts, with highest levels in the PC zone progressively decreasing toward the PP zone (Figure 3B and Figure 3-figure supplement 1A). Immunohistochemical staining for Cyp1a2 confirmed this gradient at the protein level under homeostatic conditions (0 h), with strong pericentral signal diminishing toward the portal tracts. However, following APAP treatment (3–24 h), Cyp1a2 expression was largely lost within the necrotic CV region, while the adjacent mid-zone remained the primary area retaining detectable Cyp1a2 signal (Figure 3C and 3D). This spatial shift suggests that selective necrosis of Cyp-rich pericentral hepatocytes effectively relocates the residual metabolic capacity to the surviving mid-zone parenchyma (Figure 3E).

Removed in the revised version: Next, we aimed to elucidate the cause for the stress response observed in the Mid-zone during early APAP-induced liver injury. APAP is metabolized by cytochrome P450 enzymes (Cyp family) into reactive metabolites, particularly N-acetyl-p-benzoquinone imine (NAPQI), which can induce oxidative stress and damage cellular components if not promptly detoxified by glutathione (GSH) (Figure 3A). We analyzed the gene expression levels of Cyp family enzymes in hepatocyte zones (PP, Mid, and PC) at 0, 3, 6, and 24 hours post-APAP. Our analysis revealed a gradient decrease in the expression levels of most Cyp enzymes, showing a trend from high expression in PC to lower expression in PP (Figure 3B). Scatter plots of representative Cyp enzymes, such as Cyp2e1 and Cyp1a2, clearly showed a temporal and spatial expression pattern consistent with this gradient decrease from CV to PV (Figure S2A). Immunohistochemical staining of Cyp1a2 confirmed its gradient decrease in protein expression from the CV to the PV under homeostatic conditions (0 hour). However, following APAP-induced toxicity, Cyp1a2 expression declined around the CV—where necrosis was prominent—during the injury initiation phase (3–24 hours post-APAP). Consequently, the Mid-zone adjacent to necrotic regions emerged as the primary area retaining Cyp1a2 expression (Figure 3C, D). These results suggest that selective loss of Cyp- expressing pericentral hepatocytes leads to the mid- zone becoming the primary site of residual Cyp activity (Figure 3E).

Zonal metabolism of APAP leads to stress response in Mid-zone hepatocytes.

Earlier version: (A) Schematic figure illustrating the metabolic process of liver injury caused by APAP. (B) Dot plot displays the scaled gene expression levels of the Cyp family in hepatocyte zones (zones PP, Mid, and PC) at 0, 3, 6, 24 h post-APAP. (C) Immunohistochemical staining of Cyp1a2 to detect its dynamic changes in hepatocyte zones (from CV to PV) at 0, 3, 6, 24 h post-APAP. Injured area is outline by black dashed lines. (D) Quantification of zonal distribution of Cyp1a2-positive cells in liver sections at 0, 3, 6, and 24 h post-APAP, respectively. (E) Schematic figure illustrating expression changes of Cyp1a2 before and after APAP. (F) Immunofluorescence staining of Chop protein (green) to observe the dynamic changes in its expression across hepatocyte zones (from CV to PV) at 6 h post-APAP with doses of 100, 300, or 500 mg/kg. Cyp1a2 protein (red) staining highlights the area around the CV. Cell nuclei are stained with DAPI (blue). (G) Quantification of zonal distribution of Chop-positive cells in liver sections from mice treated with varying doses of APAP– 100, 300, or 500 mg/kg. (H) Schematic explaining why mid-lobular hepatocytes display a distinct gene expression program early after APAP injury. Cyp enzymes are expressed in a decreasing gradient from the central vein (CV) to the portal vein (PV). APAP is therefore metabolized primarily by pericentral Cyp-expressing hepatocytes; their necrosis shifts the main site of residual Cyp activity to the mid-lobular zone. High local APAP metabolism then drives a specific stress-response program, accounting for the unique early transcriptional signature of the mid-zone. Data are represented as means ± SD; One-way ANOVA (D,G). *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Revised version: (A) Schematic figure illustrating the metabolic process of liver injury caused by APAP. (B) Dot plot displays the log-normalized absolute transcript abundance of the Cyp family in hepatocyte zones (zones PP, Mid, and PC) at 0, 3, 6, 24 h post-APAP. (C) Immunohistochemical staining of Cyp1a2 to detect its dynamic changes in hepatocyte zones (from CV to PV) at 0, 3, 6, 24 h post-APAP. Injured area is outline by black dashed lines. (D) Quantification of zonal distribution of Cyp1a2-positive cells in liver sections at 0, 3, 6, and 24 h post-APAP, respectively. (E) Schematic figure illustrating expression changes of Cyp1a2 before and after APAP. (F) Immunofluorescence staining of Chop protein (green) to observe the dynamic changes in its expression across hepatocyte zones (from CV to PV) at 6 h post-APAP with doses of 100, 300, or 500 mg/kg. Cyp1a2 protein (red) staining highlights the area around the CV. Cell nuclei are stained with DAPI (blue). (G) Quantification of zonal distribution of Chop-positive cells in liver sections from mice treated with varying doses of APAP– 100, 300, or 500 mg/kg. (H) Schematic explaining why mid-lobular hepatocytes display a distinct gene expression program early after APAP injury. Cyp enzymes are expressed in a decreasing gradient from the central vein (CV) to the portal vein (PV). APAP is therefore metabolized primarily by pericentral Cyp-expressing hepatocytes; their necrosis shifts the main site of residual Cyp activity to the mid-lobular zone. High local APAP metabolism then drives a specific stress-response program, accounting for the unique early transcriptional signature of the mid-zone. Data are represented as means ± SD; One-way ANOVA (D,G). *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Added in the revised version: To test whether the extent of necrosis dictates the magnitude and localization of the mid-zone stress response, we administered escalating APAP doses (100, 300, and 500 mg/kg). H&E staining confirmed a dose-dependent expansion of necrotic areas (Figure 3-figure supplement 1B--C). We then assessed Chop (the protein product of Ddit3), a key mid-zone stress marker whose transcript was upregulated at 6 h post-APAP (Figure 2B). Chop immunoreactivity was absent at the lowest dose, where necrosis was minimal. As the dose increased and the necrotic area expanded, Chop expression emerged immediately adjacent to the necrotic front and progressively shifted toward the portal veins, always surrounding the injured territory (Figure 3F and 3G). Thus, CV necrosis establishes the adjacent peri-necrotic region as the functional epicenter of residual APAP metabolism, directly linking zonal Cyp distribution to localized stress response (Figure 3H).

Removed in the revised version: To investigate this hypothesis, we administered mice with varying doses of APAP (100, 300, and 500 mg/kg) to induce different extents of liver injury. H&E staining revealed a dose-dependent increase in the size of the injured areas, with higher doses resulting in larger areas of injury (Figure S2B and S2C). We then performed immunofluorescence staining for Chop, the protein product of Ddit3, which was differentially upregulated in the Mid-zone at 6 hours post-APAP compared to 0 hours (as shown in Figure 2B), to assess the Mid-zone response. Our findings indicate that Chop expression is absent at low doses of APAP when necrosis does not occur. However, as the dose of APAP increases, the injured area expands, and Chop expression shifts toward the PV, consistently surrounding the injured area (Figure 3F and G). Together, these findings demonstrate that CV necrosis establishes the adjacent peri-necrotic region as the dominant site of APAP metabolism, contingent upon both Cyp family expression levels and APAP dosage. This spatially restricted metabolic adaptation drives subsequent region-specific transcriptional changes in the surrounding parenchyma (Figure 3H).

The Atf4-Chop axis emerges as pivotal in Mid-zone during early acute injury

Added in the revised version: To identify upstream drivers of the mid-zonal stress program, we constructed gene regulatory networks (GRNs) from the ST data and ranked the top 10 regulons per zone at 3 and 6 hours post-APAP (Figure 4A and Figure 4-figure supplement 1A). GRN analysis revealed simultaneous activation of several transcription factors (TFs) in the mid-zone across both time points, including Atf4, Fos, Nfil3, Egr1, Maff, Atf3, and Ddit3. These TFs include known stress-response regulators (Atf4, Atf3, Ddit3), immediate-early genes (Fos, Egr1), and oxidative/metabolic regulators (Maff, Nfil3), suggesting a multi-faceted adaptive program (Ohoka, Yoshii et al. 2005, Woo, Cui et al. 2009, von Scheidt, Zhao et al. 2021, Bunch, Kim et al. 2023, Lin, Hsu et al. 2024). Among the top 10 mid-zone regulons, Ddit3 exhibited the highest inferred activity, while Atf4 ranked seventh (Figure 4B and Figure 4-figure supplement 1B). Since the Atf4-Ddit3 axis is a well-established regulator of stress adaptation (Zhang, Guo et al. 2022), the coordinated activation of this transcriptional network prompted us to investigate this pathway further. Violin plots confirmed prominent mid-zone enrichment of both Atf4 and Ddit3 at 3 and 6 hours (Figure 4C and Figure 4-figure supplement 1C). GRNs illustrated connectivity between Atf4 and its targets (e.g., Gdf15, Hspa8, Ddit3) as well as between Ddit3 and its targets (e.g., Hspa8, Hmox1, Osgin1, Gadd45a, Egr1) (Figure 4-figure supplement 1D). Heatmaps further demonstrated that the target genes of both TFs were selectively expressed in mid-zone hepatocytes at 6 hours post-APAP (Figure 4D), and many of these targets overlapped with the mid-zone DEGs identified earlier. This coordinated enrichment and target overlap strongly implicate the Atf4-Ddit3 axis as a principal driver of the mid-zone transcriptional reprogramming.

Removed in the revised version: To explore master regulator of stress adaptation in Mid-zone, we constructed gene regulatory networks from the ST data. We identified the top 10 regulons in each zone at 3 and 6 hours post-APAP (Figure 4A and Figure S3A). A closer examination of the transcription factors (TFs) in the Mid-zone revealed that Atf4, Fos, Nfil3, Egr1, Maff, Atf3, and Ddit3 co-activated at both time points. Atf4, Chop, and Atf3 play crucial roles in managing cellular stress responses (Ohoka et al., 2005; Woo et al., 2009). Fos and Egr1 are immediate early genes that are rapidly activated in response to stimuli, regulating genes related to growth and stress responses (Bunch et al., 2023). Maff and Nfil3 contribute to the management of oxidative stress and metabolic processes (Lin et al., 2024; von Scheidt et al., 2021). The elevated activity of these TFs in the Mid-zone indicates that this region is undergoing a complex response to stress, serving as a critical area for cellular adaptation, survival, and potential necrosis during early liver injury.

The Atf4-Ddit3 axis emerges as pivotal in Mid-zone during early acute injury.

(A) Heatmap displays the area under the curve (AUC) scores of transcription factor (TF) motifs, estimated per dot by Single-Cell Regulatory Network Inference and Clustering (SCENIC), highlighting gene regulatory networks and differentially activated TF motifs in hepatocyte zones (PP, Mid, PC) at 6 h post-APAP. Columns represent TF motifs, rows represent dots, and color intensity indicates AUC scores. (B) Heatmap shows inferred transcription factors (TFs) activity across different zonal regions at 6 h post-APAP. Activity was quantified as regulon enrichment scores using AUCell. The color scale represents regulon activity scores. (C) Violin plots shows regulon activity of Atf4 and Ddit3 in each zonation at 6 h post-APAP. Activity was quantified as regulon enrichment scores using AUCell. (D) GO pathway analysis reveals the top 15 enriched pathways for genes co-regulated by Atf4 and Ddit3 in mid hepatocyte zones at 6 h post-APAP. (E) Heatmap showing the target genes expression of Atf4 (left) and Ddit3 (right) in hepatocyte zones (zones PP, Mid, and PC) at 6h post-APAP. (F) Immunohistochemistry staining of p-eIF2A in liver sections from mice at 0 h and 3 h post-APAP. p-eIF2A-positive hepatocytes are indicated by red arrows. The number of p-eIF2A-positive cells per field of view (FOV) is quantified. n=3 mice/group. (G) Heatmap displays stress response and ER stress-related genes from differentially expressed genes (DEGs) identified in Mid hepatocyte zones at 0 and 6 h post-APAP. Data are represented as means ± SD; Unpaired Student’s t-test (F). *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Added in the revised version: GO analysis of Atf4–Ddit3 co-regulated targets revealed enrichment for pathways related to unfolded protein response, cellular stress, autophagy, and apoptotic signaling (Figure 4E and Figure 4-figure supplement 1E--F). Immunohistochemical staining confirmed increased phosphorylation of eIF2α (p-eIF2α) in the mid-zone at 3 hours post-APAP, indicating ISR activation (Figure 4F).However, when we examined canonical ER-stress markers, we found that genes such as Eif2ak3, Atf6, and Xbp1 either remained unchanged or were downregulated at 3 and 6 hours post-APAP, despite strong Atf4 and Ddit3 upregulation (Figure 4G and Figure 4-figure supplement 1G). These data indicate that the mid-zone stress response engages the ISR through eIF2α phosphorylation but does not constitute a classical, full-scale ER stress/unfolded protein response. Instead, the response appears selectively routed through the Atf4–Ddit3 branch.

Added in the revised version: Reanalysis of a public APAP-injury scRNA-seq dataset (Matchett, Wilson-Kanamori et al. 2024). confirmed zonal hepatocyte clustering (PP, Mid, PC) at 3 and 6 hours post-APAP, with mid-zone enrichment of Hmox1, Atf3, and Ubc (Figure 4-figure supplement 2A--C). GO analysis of mid-zone upregulated genes highlighted stress-response and cell-death processes (Figure 4-figure supplement 2D--E). SCENIC analysis independently identified activation of Ddit3 and Atf3 motifs in the mid-zone (Figure 4-figure supplement 2F--G).This external dataset validates our ST-based observations and reinforces the specificity of the mid-zone stress-regulatory program.

Added in the revised version: Analysis of human snRNA-seq data from healthy and APAP-induced acute liver failure (ALF) explants (GSE223561) revealed distinct periportal (ALB, CYP2F2), mid-zonal (IGFBP2, HAMP), and pericentral (GLUL, CYP2E1) hepatocyte populations (Figure 4-figure supplement 3A-D). In healthy human livers, mid-zonal hepatocytes showed detectable expression of stress regulators (DDIT3, ATF4, HMOX1) alongside low cell-cycle gene expression. In APAP patients, mid-zonal stress regulators were further elevated; however, in contrast to our early murine model, proliferation-associated genes (MKI67, CCNB1, CDK1) were also upregulated in the same mid-zonal population (Figure 4-figure supplement 3E). Human ST data from two APAP patients revealed inter-patient heterogeneity: one specimen showed stress-gene enrichment with reduced proliferation markers in the peri-necrotic mid-zone, while the other displayed a predominantly proliferative transcriptional profile with weaker stress signaling (Figure 4-figure supplement 3F). These human data suggest that while the mid-zonal stress program is conserved, the balance between stress adaptation and proliferation is dynamic and likely stage-dependent, consistent with the temporal progression observed in our mouse model.

Removed in the revised version: Among the top 10 TFs with highest activity in the Mid-zone during early APAP-induced liver injury, Ddit3 and Atf3 were the two most highly expressed, with Atf4 ranking seventh (Figure 4B and Figure S3B). Since Ddit3 expression is often regulated by Atf4, and the Atf4-Ddit3 axis plays a critical role in stress adaptation (Zhang et al., 2022), we chose to focus on this pathway for further investigation. The violin plot analysis demonstrated that Atf4 and Ddit3 were prominently active in Mid-zone, particularly at 3 and 6 hours post-APAP (Figure 4C and Figure S3C). Gene regulatory networks illustrated associations between Atf4 and its target genes (such as Gdf15, Hspa8, and Ddit3) and Ddit3 with its targets (such as Hspa8, Hmox1, Osgin1, Gadd45a, and Egr1) at 3 hours post-APAP (Figure S3D). Heatmap further clearly showed that the target genes of Atf4 and Ddit3 are exclusively expressed in the Mid-zone hepatocytes at 6 hours post-APAP (Figure 4D). Many of these target genes were DEGs in Mid-zone during early APAP-induced liver injury, underscoring the critical role of the Atf4-Ddit3 axis in orchestrating gene expression changes in this region.

Removed in the revised version: From the regulatory networks, it is evident that many genes regulated by Atf4 and Chop overlap. Therefore, we conducted a GO pathway analysis to identify the pathways co-regulated by Atf4 and Chop. The results revealed the top 15 enriched pathways for genes co-regulated by Atf4 and Chop in Mid-zone at 3 and 6 hours post-APAP (Figure 4E and Figure S3E). These pathways encompass responses to unfolded proteins, stress, autophagy, and regulation of apoptotic signaling pathways. Additionally, GO pathway analysis identified the top 15 enriched pathways for up-regulated genes specifically in Mid-zone at 3 hours post-APAP compared to other zones, further highlighting pathways co-regulated by Atf4 and Chop in response to APAP (Figure S3F). The pathways co-regulated by Atf4 and Chop point to the activation of ISR. To confirm the occurrence of ISR in Mid-zone during the initiation stage of APAP-induced liver injury, we performed immunohistochemical staining for phosphorylated eIF2α (p-eIF2α) at 0 and 3 hours post-APAP and observed increased expression of p-eIF2α in the mid-zone (Figure 4F). While the ISR is a generalized pathway integrating endoplasmic reticulum (ER) stress and other insults through eIF2α kinases, ER stress is a subset of cellular stresses that activates the unfolded protein response (UPR) via Perk/Ire1/Atf6. Their overlap occurs at Perk-eIF2α-Atf4, but ER stress uniquely engages Ire1/Xbp1 and Atf6 for ER-specific repair (Costa-Mattioli and Walter, 2020). To evaluate the involvement of ER stress in Mid-zone, we further analyzed the expression levels of genes related to stress response and ER stress based on DEGs identified in Mid-zone at 0, 3, and 6 hours post-APAP treatment (Figure 4G and Figure S3G). Most of stress response-related genes are highly upregulated (Figure 4G and Figure S3G). However, despite the high expression of Atf4 and Ddit3, other ER stress-related genes, such as Eif2ak3, Atf6, and Xbp1 (Zhang et al., 2022) either remained unchanged or were downregulated at 3 and 6 hours compared to 0 hour post-APAP (Figure 4G and Figure S3G). This analysis highlights the unique role of the Atf4-Ddit3 axis in Mid-zone during early APAP-induced liver injury.

Removed in the revised version: To validate Mid-zone responses and transcriptional changes during early APAP-induced liver injury, we analyzed an additional scRNA-seq dataset from APAP-induced liver injury mice (Matchett et al., 2024). UMAP visualization demonstrated that hepatocytes could be clustered into three distinct zones (PP, PC, and Mid) at 3 and 6 hours post-APAP (Figures S4A and S4B). The zonal gene expression patterns were confirmed, with PC-specific genes (Glul, Cyp2e1, Cyp2a5), PP-specific genes (Alb, Cyp2f2, Sds), and Mid-zone genes (Hmox1, Atf3, Ubc) exhibiting distinct spatial distributions (Figure S4C). GO pathway analysis revealed the top 10 enriched pathways for upregulated genes in Mid-zonal hepatocytes at 3 and 6 hours post-APAP, highlighting key stress-response and cell death processes (Figures S4D and S4E). Consistent with our findings, SCENIC analysis also identified the activation of similar TF motifs across zonal regions, with distinct Ddit3 and Atf3 motifs activated in Mid-zone hepatocytes at 3 and 6 hours post-APAP (Figures S4F and S4G). These results reinforce the regulatory changes observed in our ST data, further emphasizing the unique role of mid-zonal hepatocytes in stress adaptation.

The Atf4-Ddit3 axis protects hepatocytes from liver injury

Added in the revised version: To assess the functional relevance of the Atf4–Ddit3 axis, we first examined endogenous Atf4 protein dynamics. Immunohistochemical staining revealed progressive nuclear Atf4 accumulation in mid-zone hepatocytes between 0 and 6 hours post-APAP (Figure 5A). This temporal and spatial pattern aligns with the transcriptional upregulation observed in our ST data, confirming mid-zone-specific Atf4 activation.

Removed in the revised version: To investigate the role of Atf4-Ddit3 axis in cell fate determination, we first examined Atf4 expression in liver sections by immunohistochemical staining at 0 and 6 hours post-APAP. The results showed gradual increased expression of Atf4 in Mid-zone during early APAP-induced liver injury (Figure 5A). Subsequently, we employed the adeno-associated virus 8 (AAV8) and thyroxine-binding globulin (TBG) promoter method to overexpress Atf4 specifically in hepatocytes (Figure 5B). Atf4 overexpression significantly reduced APAP-induced liver injury, as evidenced by H&E staining and serum ALT levels (Figure 5C and 5D). Furthermore, immunohistochemical staining of Atf4 indicated its localization in hepatocyte nuclei around the injured area (Figure 5E). Immunofluorescence and immunohistochemical staining for Ki67 and p-eIF2α, respectively, showed that Atf4 overexpression notably decreased Ki67 expression and increased p-eIF2α expression (Figure 5F and 5G). Together, these findings suggest that elevated Atf4 expression enhances ISR, offering stress protection under conditions of reduced proliferation.

The Atf4-Ddit3 axis protects hepatocytes from liver injury.

(A) Immunohistochemical staining of Atf4 in liver sections at 0 and 6 h post-APAP. Red arrows indicate Atf4-positive hepatocytes. Zonal distribution of Atf4-positive cells in liver sections at 6 h post-APAP is quantified. The statistic is the percentage of Atf4-positive hepatocytes in each layer over the total number of Atf4-positive hepatocytes. n=3 mice. (B) Schematic figure illustrating the overexpression of Atf4 via AAV in hepatocytes of wildtype C57BL/6J mice. Overexpression of EGFP is used as a control. (C) H&E staining showing liver morphology from mice overexpressing AAV-TBG-EGFP or AAV-TBG-Atf4 at 6 h post-APAP. The injured area is outlined by black dashed lines. The percentage of injured area is quantified. n=3 mice/group. Necrotic areas outlined by loss of cellular architecture on H&E. (D) Serum levels of ALT is measured in mice overexpressing AAV-TBG-EGFP or AAV-TBG-Atf4 at 6, 24, 48, and 72 h post-APAP. n=6 mice/group. (E) Immunohistochemistry staining of Atf4 in liver sections from mice overexpressing AAV-TBG-EGFP or AAV-TBG-Atf4 at 24 h post-APAP. Atf4-positive hepatocytes are indicated by red arrows. The number of Atf4-positive cells per field of view (FOV) is quantified. n=9 mice/group. (F) Immunofluorescence staining of Ki67 (red) was performed to assess proliferating hepatocytes in mice overexpressing AAV-TBG-EGFP or AAV-TBG-Atf4 at 72 h post-APAP. Cell nuclei were stained with DAPI (blue). The number of Ki67-positive cells per FOV is quantified. n=6 mice/group. (G) Immunohistochemistry staining of p-eIF2A in liver sections from mice overexpressing AAV-TBG-EGFP or AAV-TBG-Atf4 at 24 h post-APAP. p-eIF2A-positive hepatocytes are indicated by red arrows. The number of p-eIF2A-positive cells per FOV is quantified. n=6 mice/group. Data are represented as means ± SD; Unpaired Student’s t-test (C-G). One-way ANOVA (A). *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Added in the revised version: To test causality, we overexpressed Atf4 specifically in hepatocytes using adeno-associated virus 8 (AAV8) - thyroxine-binding globulin (TBG) (Figure 5B). Atf4 overexpression markedly attenuated APAP-induced liver injury, as evidenced by reduced necrotic areas on H&E staining and significantly lower serum ALT levels (Figure 5C and 5D). Immunohistochemical staining confirmed nuclear Atf4 localization in hepatocytes adjacent to injured regions (Figure 5E). Furthermore, Atf4 overexpression resulted in decreased Ki67⁺proliferating cells and increased p-eIF2α immunoreactivity (Figure 5F and 5G). Together, these data demonstrate that elevated Atf4 actively promotes hepatocyte survival by enhancing ISR signaling, suppressing proliferation, and mitigating early liver damage.

The Atf4-Ddit3 axis protects the liver by pausing proliferation via Btg2

Earlier version: Atf4 often works with Ddit3 to mediate stress response, particularly in the context of the ISR (Kaspar et al., 2021; Ohoka et al., 2005; Woo et al., 2009). Given the low expression of Atf4 and high expression of Ddit3 in the Mid-zone, we performed Cut&Run to identify the genomic regions where Ddit3 binds during the initiation phases of APAP-induced liver injury. This approach was used to investigate how the Atf4-Ddit3 axis regulates the proliferation of Mid-zone hepatocytes. Compared with the input control, we observed heightened Ddit3-binding signals within the chromatin domains (Figure 6A). We further overlapped 6157 genes where Ddit3 binds, as identified by Cut&Run, with 47 DEGs specifically in the Mid-zone hepatocytes, identifying 39 overlapping genes termed Ddit3-binding DEGs (Figure 6B). Integrative Genomics Viewer (IGV) plots illustrate Ddit3-binding peaks and input controls at specific genomic sites such as Btg2, Atf3, Dnajb1, and Egr1 regulatory elements (Figure 6C).

Revised version: Atf4 and Ddit3 are key transcriptional regulators of ISR, and their downstream target genes converge on overlapping stress-responsive pathways (Ohoka, Yoshii et al. 2005, Woo, Cui et al. 2009, Kaspar, Oertlin et al. 2021). To dissect the downstream effectors of the Atf4–Ddit3 axis, we mapped Ddit3 genomic occupancy by Cut&Run at 6 h post-APAP—the peak of mid-zone Ddit3 activity during the injury-initiation phase. Compared with input control, we detected robust Ddit3-binding signals across chromatin domains (Figure 6A). Overlapping 6,157 Ddit3-bound genes with the 47 mid-zone–specific DEGs identified 39 overlapping targets, termed Ddit3-binding DEGs (Figure 6B). IGV tracks confirmed Ddit3 occupancy at regulatory elements of Btg2, Atf3, Dnajb1, and Egr1 (Figure 6C).

The Atf4-Ddit3 axis protects the liver by pausing proliferation via Btg2.

(A) Heatmap of Ddit3-binding signals in chromatin domains. (B) Overlap of Ddit3-binding genes by Cut&Run with differentially expressed genes (DEGs) in mid zone was shown. These genes were referred as Ddit3-binding DEGs. (C) Integrative Genomics Viewer (IGV) plot shows the Cut&Run peaks of Ddit3 and input on Btg2-site, Ifrd1-site, Fos-site and Egr1-site. (D) Immunohistochemical staining of Btg2 in liver sections at 0 and 6 h post-APAP. Red arrows indicate Btg2-positive hepatocytes. (E) Schematic figure illustrating overexpression of Ddit3 via AAV in hepatocytes of wildtype C57BL/6J mice. Overexpression of EGFP is used as a control. (F) Immunohistochemical staining of Btg2 in liver sections from mice overexpressing either AAV-TBG-EGFP or AAV-TBG-Ddit3 at 12 h post-APAP. Btg2-positive hepatocytes are indicated by red arrows. Btg2-positive hepatocytes in liver sections is quantified. n=4 mice/group. (G) Schematic figure illustrating overexpression of Btg2 via AAV in hepatocytes of wildtype C57BL/6J mice. Overexpression of EGFP is used as a control. (H) H&E staining shows morphology of livers from AAV-TBG-EGFP or AAV-TBG-Btg2 overexpressing mice at 6 h post-APAP. Injured area is outlined by black dashed lines. The percentage of injury area is quantified. n=3 mice/group. Necrotic areas outlined by loss of cellular architecture on H&E. (I) Serum levels of ALT is measured at 6h post-APAP in AAV-TBG-EGFP or AAV-TBG-Btg2 overexpressing mice. n=3 mice/group. (J) Immunohistochemical staining of Ki67 in liver sections from mice overexpressing either AAV-TBG-EGFP or AAV-TBG-Btg2 at 72 h post-APAP. Ki67-positive hepatocytes are indicated by red arrows. Ki67-positive cells in liver sections is quantified. n=4 mice/group. (K) Schematic figure illustrating knockdown of Btg2 via AAV in hepatocytes of wildtype C57BL/6J mice. IP479-EFS-sgEGFP is used as a control. (L) Serum levels of ALT is measured at 6h post-APAP in IP479-EFS-sgEGFP or IP479-EFS-sgBtg2 knockdown mice. n=3 mice/group. (M) H&E staining shows morphology of livers from IP479-EFS-sgEGFP or IP479-EFS-sgBtg2 knockdown mice at 6 h post-APAP. Injured area is outlined by black dashed lines. The percentage of injury area is quantified. n=3 mice/group. Necrotic areas outlined by loss of cellular architecture on H&E. (N) Immunohistochemical staining of Ki67 in liver sections from IP479-EFS-sgEGFP or IP479-EFS-sgBtg2 knockdown mice at 72 h post-APAP. Ki67-positive cells in liver sections is quantified. n=3 mice/group. Data are represented as means ± SD; Unpaired Student’s t-test (F, H-N). **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Earlier version: Among the Ddit3-binding DEGs, we observed B-cell translocation gene 2 (Btg2), which is cell cycle inhibitor (Park et al., 2008; Stupfler et al., 2016; Yuniati et al., 2019). Moreover, Btg2 is one of the target genes that coregulated by Ddit3 and Atf4. Immunohistochemical staining of Btg2 in liver sections at 0 and 6 hours post-APAP revealed increased expression of Btg2 around the injured area in response to APAP (Figure 6D). Additionally, we investigated whether Ddit3 could upregulate Btg2 expression in Mid-zone. We overexpressed Ddit3 in the hepatocytes of wildtype C57BL/6J mice using the AAV8 method, with EGFP overexpression serving as a control (Figure 6E). Immunohistochemical staining of Btg2 in liver sections from mice overexpressing either AAV-TBG-EGFP or AAV-TBG-Ddit3 at 12 hours post-APAP indicated that Ddit3 overexpression significantly upregulated Btg2 expression post-APAP (Figure 6F). This result suggests that Btg2 is a downstream target of Ddit3.

Revised version: Among the Ddit3-binding DEGs, B-cell translocation gene 2 (Btg2) is cell cycle inhibitor (Park, Kim et al. 2008, Stupfler, Birck et al. 2016, Yuniati, Scheijen et al. 2019). Moreover, Btg2 is also coregulated by Ddit3 and Atf4. Immunohistochemical staining further revealed increased Btg2 protein expression around the injured area at 6 h post-APAP (Figure 6D). To determine whether Ddit3 drives Btg2 expression, we overexpressed Ddit3 in hepatocytes via AAV8-TBG (Figure 6E). Ddit3 overexpression significantly elevated Btg2 protein levels at 12 h post-APAP compared to EGFP controls (Figure 6F). We next overexpressed Btg2 directly (Figure 6G). Btg2 overexpression significantly reduced APAP-induced injury, as shown by decreased necrotic areas and serum ALT levels (Figure 6H, 6I), and markedly reduced Ki67⁺ hepatocyte numbers (Figure 6J).

Added in the revised version: To rule out potential confounds associated with AAV transduction, we measured Cyp2e1 protein levels, which showed only a modest decrease in both Atf4- and Btg2-overexpressing groups—well below the threshold known to confer significant protection (Figure 6-figure supplement 1A) (Ganetsky, Berg et al. 2019, Ye, Chen et al. 2022). AAV transduction efficiencies averaged 32% (EGFP), 18% (Atf4), and 23% (Btg2), with individual efficiency negatively correlating with ALT levels (e.g., Pearson r = –0.7681, p = 0.0260 for Atf4; Figure 6-figure supplement 1B--C). These controls confirm that the protective effects are attributable to transgene expression rather than nonspecific Cyp2e1 inhibition. Forced Btg2 expression phenocopies the protective, proliferation-suppressive state observed during early injury, positioning Btg2 as a critical downstream mediator of the stress response.

Added in the revised version: To determine whether Btg2 is necessary for the protective arrest, we knocked down Btg2 using the AAV8-CasRx-sgRNA system (Figure 6K) (Zhou, Su et al. 2020). Despite a moderate (∼30%) reduction in Btg2 expression (Figure 6-figure supplement 1D), knockdown resulted in a significant ∼2-fold increase in serum ALT, ∼1.5-fold expansion of necrotic areas, and ∼1.8-fold increase in Ki67⁺hepatocytes compared to sgEGFP controls (Figure 6L--N). These data suggest that loss of Btg2 abrogates the protective brake on proliferation, sensitizing the liver to injury and permitting unscheduled cell-cycle activity under stress conditions.

Added in the revised version: To assess the broader applicability of this mechanism, we tested a carbon tetrachloride (CCl4)-induced acute liver injury model, which causes centrilobular necrosis through free-radical pathways rather than cytochrome P450–dependent metabolite formation (Wang, Wang et al. 2024). At 18 h post-CCl₄ injection, we observed elevated p-eIF2α, Atf4, Chop, and Btg2 levels alongside reduced Ki67 expression near injury sites (Figure 6-figure supplement 2A--G). This conservation across mechanistically distinct toxicants suggests that ISR-mediated Btg2 induction and subsequent proliferation arrest represent a general adaptive response to centrilobular injury. Collectively, these data establish a conserved functional cascade in which the Atf4-Ddit3 axis induces Btg2 to enforce transient proliferative restraint, prioritizing survival during early acute liver injury and preserving the capacity for subsequent regeneration.

Removed in the revised version: To explore the role of Btg2 in liver injury and hepatocyte proliferation, we overexpressed Btg2 in hepatocytes of wild-type C57BL/6J mice (Figure 6G). H&E staining and serum ALT levels indicated that Btg2 overexpression significantly reduced APAP-induced liver injury (Figure 6H and 6I). Ki67 staining further demonstrated that Btg2 overexpression significantly decreased hepatocyte proliferation (Figure 6J). To assess potential confounds from AAV-based overexpression, we performed western blot analysis of Cyp2e1 protein levels in AAV-EGFP control mice and mice overexpressing Atf4 or Btg2 without APAP treatment. Compared with the AAV-EGFP group, Cyp2e1 expression was slightly decreased in the AAV-Atf4 and AAV-Btg2 groups (Figure S5A). Importantly, this reduction is well below the threshold of Cyp inhibition known to confer significant protection against APAP-induced liver injury (Ganetsky et al., 2019; Ye et al., 2022a). Therefore, the modest decrease in Cyp2e1 expression observed here is unlikely to account for the pronounced protective effects of Atf4 or Btg2 overexpression. We also quantified the penetrance of AAV-mediated overexpression. Mice were injected with 1.2 × 10¹¹ viral genome copies per animal, yielding average transduction efficiencies of 32%, 18%, and 23% for AAV-EGFP, AAV-Atf4, and AAV-Btg2, respectively (Figure S5B). Individual animal transduction efficiency showed a negative correlation with serum ALT levels (e.g. Pearson r = –0.7681, p = 0.0260 for Atf4; Figure S5C), demonstrating that greater transgene expression associates with stronger protection. Moreover, we used the CasRx-sgRNA system (Zhou et al., 2020) to knock down Btg2 by overexpressing either IP479-EFS-sgEGFP or IP479-EFS-sgBtg2 (Figure 6K). Knockdown of Btg2 using AAV8- CasRx achieved a moderate (∼30%) reduction in Btg2 expression (Figure S5D). Despite this partial efficiency, we observed a significant increase in serum ALT levels (∼2- fold) (Figure 6L), expansion of necrotic areas (∼1.5- fold) (Figure 6M), and a marked increase in Ki67⁺ hepatocytes (∼1.8- fold) compared to control mice (Figure 6N). Together, these findings illustrate Atf4-Ddit3 axis inhibits hepatocytes proliferation in Mid-zone via Btg2 during early APAP-induced liver injury.

Removed in the revised version: To test whether the proposed mechanism applies to other liver injuries, we employed mouse models of partial hepatectomy (PHx) and carbon tetrachloride (CCl4)-induced acute liver injury. CCl4 causes centrilobular necrosis (similar to acetaminophen) but induces damage via free radical–mediated mechanisms rather than through cytochrome P450–dependent reactive metabolite formation (Wang et al., 2024). In our CCl4 model (administered intraperitoneally in corn oil, with samples collected 18 h post- injection), the ISR was activated around injury sites, accompanied by decreased proliferation, as evidenced by increased expression of p- eIF2α, Atf4, Chop, and Btg2, along with reduced Ki67 expression (Figure S6A–G). Partial hepatectomy involves mechanical resection, resulting in ischemic injury in the remaining portion of the left lobe due to compromised blood supply (Qu et al., 2022). In this model (examined 24 h after surgery), ISR activation was similarly observed around ischemic injury sites, with increased expression of p- eIF2α, Atf4, Chop, and Btg2 and undetectable Ki67 expression (Figure S7A–G). Together, these findings suggest that the proposed mechanism may be broadly applicable to other types of liver injury.

Discussion

Earlier version: Our study reveals that during the early phase of APAP-induced liver injury, mid-zone hepatocytes undergo a transient proliferation arrest, orchestrated by a distinctive transcriptional program centered on the Atf4-Chop axis. Spatial transcriptomics and functional assays demonstrate that this adaptive response promotes survival by upregulating the cell cycle inhibitor Btg2, highlighting a trade-off between cytoprotection and proliferative capacity. Notably, zonal metabolic heterogeneity in APAP processing drives preferential activation of these stress pathways in mid-zone hepatocytes. These findings uncover a critical mechanism by which hepatocytes prioritize stress adaptation before initiating regeneration, offering new insights into the spatiotemporal regulation of liver repair (Figure 7).

Revised version: Our study supports a model that during the early phase of APAP-induced liver injury, mid-zone hepatocytes undergo a transient proliferative restraint orchestrated by a distinctive transcriptional program centered on the Atf4–Chop–Btg2 axis. Spatial transcriptomics and functional assays demonstrate that this adaptive response promotes survival by upregulating the cell-cycle inhibitor Btg2, highlighting a trade-off between cytoprotection and proliferative capacity. Notably, zonal metabolic heterogeneity in APAP processing drives preferential activation of these stress pathways in mid-zone hepatocytes. These findings uncover a critical mechanism by which hepatocytes prioritize stress adaptation before initiating regeneration, offering new insights into the spatiotemporal regulation of liver repair (Figure 7).

The Atf4-Ddit3 axis mediates integrated stress protection in mid-zone hepatocytes at the expense of proliferative capacity during early AILI.

Schematic model depicting the adaptive response of mid-zone hepatocytes to acute liver injury, wherein activation of the Atf4-Ddit3 pathway upregulates Btg2 to induce transient cell cycle arrest, balancing survival and regeneration.

Earlier version: Our study reveals that mid-zone hepatocytes employ a protective proliferation arrest during early APAP-induced liver injury, mediated by the Atf4-Chop axis of the integrated stress response (ISR). While the ISR shares upstream initiators with ER stress (e.g., eIF2α phosphorylation), its outcomes are context-dependent. ER stress specifically activates the UPR through sensors inositol-requiring enzyme 1α (Ire1α), activating transcription factor 6 (Atf6), and protein kinase RNA-like ER kinase (Perk) (Wang et al., 2006), with genetic studies demonstrating their critical roles in APAP-induced liver injury knockout of Ire1α, X-box binding protein 1 (Xbp1), or Chop consistently attenuates liver damage, confirming the UPR’s contribution to hepatotoxicity (Hur et al., 2012; Nagy et al., 2007; Uzi et al., 2013; Ye et al., 2022b). Unlike classical ER stress, we find that mid-zone hepatocytes leverage the Atf4-Chop axis to transiently halt proliferation via Btg2, favoring survival. This aligns with emerging evidence that Chop can paradoxically support adaptation in mitochondrial stress (Kaspar et al., 2021; Kress et al., 2023), suggesting that its role in fine-turning of ISR in mammals.

Revised version: While the ISR shares upstream initiators with ER stress (e.g., eIF2α phosphorylation), its outcomes are context-dependent. ER stress specifically activates the UPR through Ire1α, Atf6, and Perk (Wang, Thomas et al. 2006), with genetic studies demonstrating that knockout of Ire1α, Xbp1, or Ddit3 consistently attenuates liver damage, confirming the UPR’s contribution to hepatotoxicity (Nagy, Kardon et al. 2007, Hur, So et al. 2012, Uzi, Barda et al. 2013, Ye, Chen et al. 2022). Unlike classical ER stress, we find that mid-zone hepatocytes leverage the Atf4–Chop axis—defined here as a functional pathway of co-regulated transcriptional targets—to transiently halt proliferation via Btg2, favoring survival. This aligns with emerging evidence that Ddit3 can paradoxically support adaptation in mitochondrial stress (Kaspar, Oertlin et al. 2021, Kress, Jessen et al. 2023), suggesting that Ddit3 plays a context-dependent role in fine-tuning the ISR in mammals.

Earlier version: Our findings on the early stress-induced proliferation arrest in mid-lobular hepatocytes (3–12h post-APAP) complement yet fundamentally extend prior knowledge of zonal liver regeneration. While foundational studies established perinecrotic hepatocytes as drivers of later proliferation (>24h) (Bajt et al., 2003), and recent work delineated CXCR2+ hepatocytes as pro-proliferative effectors at moderate doses (Nguyen et al., 2022) or CXCL14/p21-mediated senescence at high doses (Umbaugh et al., 2024), our work resolves an earlier, spatiotemporally distinct adaptive phase. Crucially, we identify mid-zone hepatocytes—not merely as transitional cells en route to proliferation—but as orchestrators of an active survival strategy via the Atf4-Chop-Btg2 axis, which redirects resources from division to cytoprotection (e.g., redox defense). This mechanism operates before necrosis expansion and is distinct from periportal/pericentral stress responses, which focus on 2h post-injury and lack mid-zone resolution (Umbaugh et al., 2021). By integrating spatial transcriptomics with functional validation (Cut&Run, AAV-Btg2 modulation), we establish causal links between stress signaling and quiescence—unlike descriptive profiling in earlier zonal studies (Wu et al., 2024). Thus, our model bridges early stress responses and later regeneration, revealing how mid-zone hepatocytes enable subsequent repair by prioritizing survival during the critical initiation window.

Revised version: Our work resolves an early adaptive phase (3–12 h) of proliferation arrest that is temporally and functionally distinct from previously described zonal regeneration programs, including later perinecrotic proliferation (>24 h) (Bajt, Knight et al. 2003), CXCR2⁺ pro-proliferative responses (Nguyen, Umbaugh et al. 2022), and p21-mediated senescence (Umbaugh, Nguyen et al. 2024). Rather than acting as transitional cells, mid-zone hepatocytes deploy the Atf4–Chop–Btg2 axis to transiently redirect resources from cell division toward cytoprotective pathways, such as redox defense. This spatially restricted mechanism operates before overt necrosis expansion and, unlike earlier periportal/pericentral stress profiles (Umbaugh, Ramachandran et al. 2021), is causally supported by integrated spatial transcriptomics and functional validation (Cut&Run, AAV modulation). Our findings therefore bridge early stress adaptation with subsequent liver repair, revealing that mid-zone hepatocytes may preserve regenerative capacity by prioritizing survival over proliferation during the critical injury-initiation window.

Earlier version: The antiproliferative activity of Btg2 has been studied through an integrated network of transcriptional and post-transcriptional cell cycle regulatory mechanisms. A study reveals that Btg2 functions as a molecular bridge between poly(A)-binding protein cytoplasmic 1 (Pabpc1) and the CCR4-NOT deadenylase complex subunit CNOT7 (Caf1), facilitating accelerated degradation of cell cycle-related transcripts through enhanced poly(A) tail removal (Stupfler et al., 2016). At the transcriptional level, Btg2 exerts specific control over G2/M progression by disrupting the positive feedback loop between cyclin B1 and its transcriptional activator Foxm1 (forkhead box protein M1), thereby suppressing cyclin B1 expression and inducing cell cycle arrest (Park et al., 2008). Furthermore, Btg2 demonstrates broader cell cycle inhibitory capacity through its interactions with cyclin-dependent kinases (Cdks) during G1/S transition and its ability to modulate the tumor suppressor p53 (TP53) (Yuniati et al., 2019). In the context of APAP-induced liver injury, the stress-responsive Atf4-Chop pathway induces Btg2 expression in mid-zonal hepatocytes, where it may coordinate a comprehensive proliferation arrest through these complementary mechanisms. This multilayered regulatory strategy enables hepatocytes to temporarily suspend cell division and redirect resources toward cellular repair and stress adaptation during the acute phase of toxic injury.

Revised version: The antiproliferative activity of Btg2 has been studied through an integrated network of transcriptional and post-transcriptional regulatory mechanisms. At the post-transcriptional level, Btg2 bridges poly(A)-binding protein cytoplasmic 1 (Pabpc1) with the CCR4-NOT deadenylase complex, accelerating degradation of cell cycle-related transcripts (Stupfler, Birck et al. 2016). Transcriptionally, Btg2 disrupts the cyclin B1–Foxm1 positive feedback loop, suppressing cyclin B1 expression and inducing G2/M arrest (Park, Kim et al. 2008). Btg2 also inhibits G1/S transition through cyclin-dependent kinase interactions and modulates p53 activity (Yuniati, Scheijen et al. 2019). In the context of APAP-induced liver injury, the stress-responsive Atf4–Chop pathway induces Btg2 expression in mid-zonal hepatocytes, where it likely coordinates a comprehensive proliferation arrest through these complementary mechanisms. This multilayered regulatory strategy enables hepatocytes to temporarily paused cell proliferation and redirect resources toward cellular repair and stress adaptation during the acute phase of toxic injury.

Added in the revised version: A key limitation of this study is the scarcity of early-phase APAP-induced liver injury human samples, which precludes definitive validation of the clinical relevance of these zonal adaptations. The apparent differences between the human and murine datasets are likely attributable to the stage of liver injury analyzed. Human samples were obtained from end-stage liver explants, in which regenerative programs are already active, whereas our mouse studies focused on the early (3–12 h) injury phase preceding overt regeneration. Consistent with this interpretation, one ST specimen displayed a stress-dominant mid-zonal transcriptional program resembling our murine observations, whereas the second exhibited a predominantly proliferative profile. These findings suggest that the balance between stress adaptation and hepatocyte proliferation is dynamic and strongly influenced by the temporal stage of liver injury. Although the available human datasets do not capture the early injury window examined in our mouse model, they support conservation of the mid-zonal stress-response program across species.

Earlier version: A key limitation of this study is the scarcity of early-phase APAP-induced liver injury human samples, which currently precludes definitive validation of the clinical relevance of these zonal adaptations. To address this gap, we performed a preliminary analysis of published ST data from APAP-overdose patients (Matchett et al., 2024). Intriguingly, in one of two analyzed patients, mid-zone hepatocytes exhibited transcriptional signatures remarkably consistent with our murine findings, including: (1) upregulation of Atf4-Chop pathways, and (2) downregulation of cell proliferation genes. While the small sample size (n=2) prevents definitive conclusions, these preliminary observations suggest potential conservation of stress-responsive mechanisms between murine and human mid-zone hepatocytes during APAP-induced liver injury. Further studies with larger clinical cohorts will be essential to verify these findings and establish their translational significance. While our ST analysis and Ki67 staining consistently indicated hepatocytes as the primary proliferating compartment during early APAP-induced liver injury, we recognize that non-parenchymal cells (e.g., immune cells, stellate cells) may contribute to division signals. Single-cell resolution techniques would be required to completely exclude these contributions, representing a technical limitation of bulk spatial transcriptomics. Additionally, our ST data alone do not independently confirm zonal cell cycle arrest: ST lacks overt separation in S/G2/M scores, uses non-canonical markers (e.g., Nasp, Cks1b), and has low sensitivity for rare Ki67⁺ events (∼1% of spots). Thus, our primary evidence for midlobular proliferation dynamics comes from Ki67 immunohistochemistry (Figure 1), with ST findings (Figure 2E) shown as exploratory and directionally consistent with transient impairment, but not as definitive proof of cell cycle arrest.

Revised version: We also acknowledge several technical limitations of our spatial transcriptomics approach. While ST analysis and Ki67 staining consistently identified hepatocytes as the primary proliferating compartment during early APAP-induced liver injury, non-parenchymal cells (e.g., immune cells, stellate cells) may contribute to division signals. Single-cell resolution techniques would be required to exclude these contributions fully. Additionally, our ST data alone do not independently confirm zonal proliferative restraint: ST lacks overt separation in S/G2/M scores, uses non-canonical markers (e.g., Nasp, Cks1b), and has low sensitivity for rare Ki67⁺ events (∼1% of spots). Thus, our primary evidence for mid-lobular proliferation dynamics comes from Ki67 immunohistochemistry (Figure 1), with ST findings (Figure 2E) shown as exploratory and directionally consistent with transient impairment, but not as definitive proof of proliferative restraint.

In summary, our study uncovers a spatially coordinated stress-adaptation mechanism in mid-zone hepatocytes during early APAP-induced liver injury, where the Atf4–Chop axis transiently suppresses proliferation via Btg2 upregulation. By elucidating how hepatocytes prioritize survival before initiating repair, our work provides a framework for understanding the dynamic interplay between stress adaptation and regeneration in tissues.

Materials and methods

Animal experiments and procedures

Removed in the revised version: Animals

Earlier version: C57BL/6J mice (strain no. N000013) were used as wild-type (WT) mice. All mouse colonies were maintained at the Animal Core Facility of Yunnan University. The animal studies were approved by the Yunnan University Institutional Animal Care and Use Committee (IACUC, Approval No. YNU20220387). For APAP treatment, mice (8-12 weeks old) were fasted overnight (5:00pm to 9:00am) before i.p. injected with APAP (Sigma, A7085) at a dose of 300 mg/kg for male mice, as female mice are less susceptible to APAP-induced liver injury (Guerrero Munoz and Fearon, 1984). Male mice have been the choice in the vast majority of the studies of APAP-induced liver injury reported in the literature (Ben-Moshe et al., 2022b; Campion et al., 2008; Matchett et al., 2024; Shan et al., 2021). Therefore, we used male mice in the majority of the experiments presented. In one experiment, mice were i.p. injected with various doses of APAP (100, 300, 500mg/kg). For carbon tetrachloride (CCl4) treatment, mice were i.p. injected with corn oil or CCl4 (Macklin, Cat #C805329) at a dose of 0.8ul/g (corn oil dilute to 20%). For 70% partial hepatectomy, mice were placed in a supine position and secured under isoflurane inhalation anesthesia. A midline upper abdominal incision is made. The liver is gently exposed, and the left lateral lobe and the median lobe (comprising the left and right median lobes) are mobilized. The hepatic pedicle at the base of each lobe is ligated with 6-0 silk suture. After confirming the absence of bleeding, the lobes are excised just distal to the ligature, removing approximately 70% of the liver mass. The resection site is inspected for active bleeding. The abdominal incision is closed in layers with absorbable suture. Postoperatively, the mouse is placed on a heating pad for recovery.

Revised version: Animals C57BL/6J mice (strain no. N000013) were used as wild-type (WT) mice. All mouse colonies were maintained at the Animal Core Facility of Yunnan University. The animal studies were approved by the Yunnan University Institutional Animal Care and Use Committee (IACUC, Approval No. YNU20220262). For APAP treatment, mice (8-12 weeks old) were fasted overnight (5:00pm to 9:00am) before i.p. injected with APAP (Sigma, A7085) at a dose of 300 mg/kg for male mice, as female mice are less susceptible to APAP-induced liver injury (Guerrero Munoz and Fearon 1984). Male mice have been the choice in the vast majority of the studies of APAP-induced liver injury reported in the literature (Campion, Johnson et al. 2008, Shan, Li et al. 2021, Ben-Moshe, Veg et al. 2022, Matchett, Wilson-Kanamori et al. 2024). Therefore, we used male mice in the majority of the experiments presented. In one experiment, mice were i.p. injected with various doses of APAP (100, 300, 500mg/kg). For carbon tetrachloride (CCl4) treatment, mice were i.p. injected with corn oil or CCl4 (Macklin, Cat #C805329) at a dose of 0.8ul/g (corn oil dilute to 20%).

In some experiments, APAP-treated mice were pre-injected intraperitoneally (i.p.) with either AAV-overexpression plasmid (AAV-TBG-Atf4/Ddit3/Btg2) or AAV-sgRNA knockdown plasmid (IP479-EFS-sgBtg2) for 3 weeks. AAV-TBG-EGFP and IP479-EFS-sgEGFP were used as controls, respectively. ALT measurement was performed using a diagnostic assay kit (Teco Diagnostics, Anaheim CA) to assess liver injury.

Histology, Immunohistochemistry & immunofluorescence

H&E staining

Tissues were fixed overnight at 4°C using buffered 10% paraformaldehyde (Sangon Biotech, A500684-0500) and subsequently embedded in paraffin. For H&E staining, paraffin-embedded slides were deparaffinized and rehydrated following standard protocols. The sections were briefly immersed in hematoxylin for 30 seconds, rinsed in water, and then stained with Eosin solution for 3 minutes. After staining, the sections were dehydrated in ethanol, cleared in xylene, and mounted. Slides were examined under a microscope (Scan System, SQS-1000) and analyzed using ImageJ. Necrotic areas outlined by loss of cellular architecture on H&E.

Immunofluorescence on frozen section

Earlier version: Immunofluorescence staining was conducted on frozen sections of fresh liver tissues. Initially, the tissues were fixed using 2% paraformaldehyde for 1 h and subsequently dehydrated overnight in a 30% sucrose solution. The following day, tissue embedding was performed using OCT (SAKURA, 4583), after which ultra-thin sections of 5μm were sliced. Permeabilization was achieved using 0.2% Triton X-100 for 10 minutes at room temperature, followed by blocking with 5% normal goat serum (VivaCell, C2530-0100). Primary antibodies against mouse GS (Proteintech, 11037-2-AP, 1:200), Chop (Proteintech, 15204-1-AP, 1:1000), Cyp1a2 (Proteintech, 15204-1-AP, 1:50) and Ki67 (Abcam, ab180569, 1:1000) were then applied. Secondary antibodies (Affinipure Goat Anti-Rabbit 488-conjugated, Jackson ImmunoResearch, 111-545-003, 1:400; Goat anti-Rabbit IgG(H+L) Cross Adsorbed Secondary Antibody,Alexa Fluor 568, Invitrogen, A11011, 1:1000 and Goat Anti-Mouse 594-conjugated, Jackson ImmunoResearch, 115-585-003, 1:200) were used accordingly. After washing, the slides were mounted using antifade medium, and nuclei were stained with DAPI (Beyotime, C1006, prediluted). Images were captured using an Olympus BP80 microscope and analyzed with ImageJ.

Revised version: Immunofluorescence staining was conducted on frozen sections of fresh liver tissues. Initially, the tissues were fixed using 2% paraformaldehyde for 1 h and subsequently dehydrated overnight in a 30% sucrose solution. The following day, tissue embedding was performed using OCT (SAKURA, 4583), after which ultra-thin sections of 5μm were sliced. Permeabilization was achieved using 0.2% Triton X-100 for 10 minutes at room temperature, followed by blocking with 5% normal goat serum (VivaCell, C2530-0100). Primary antibodies against mouse GS (Proteintech, 11037-2-AP, 1:200), Chop (Proteintech, 15204-1-AP, 1:1000), Cyp1a2 (Proteintech, 15204-1-AP, 1:50) and Ki67 (Abcam, ab180569, 1:1000) were then applied. Secondary antibodies (Affinipure Goat Anti-Rabbit 488-conjugated, Jackson ImmunoResearch, 111-545-003, 1:400; Goat anti-Rabbit IgG (H+L) Cross-Adsorbed Secondary Antibody, Alexa Fluor ™ 568, Invitrogen, A11011, 1:1000, Alexa Fluor 488-conjugated Affinipure Goat Anti-Mouse IgG+IgM(H+L), Jackson ImmunoResearch, 115-545-044, 1:400 and Alexa Fluor® 594 AffiniPure® Goat Anti-Mouse IgG (H+L), Jackson ImmunoResearch, 115-585-003, 1:200) were used accordingly. After washing, the slides were mounted using antifade medium, and nuclei were stained with DAPI (Beyotime, C1006, prediluted). Images were captured using an Olympus BP80 microscope and analyzed with ImageJ.

Immunohistochemistry

Livers were fixed in 10% buffered formalin for 24 hours before paraffin embedding. For immunohistochemical staining of paraffin-embedded liver tissue, 5 μm sections were prepared, deparaffinized, and subjected to antigen retrieval by microwaving for 20 minutes. Antigen retrieval was performed in pH 6.0 sodium citrate buffer (Ki-67, P62, Chop, Cyp1a2, Cyp2f2), pH 6.4 sodium citrate buffer (p-eIF2A), or pH 9.0 Tris-EDTA buffer (Atf4, Btg2). Sections were blocked with 3% hydrogen peroxide and permeabilized with 0.2% Triton X-100 for 10 minutes at room temperature, followed by blocking with 5% normal goat serum (VivaCell, C2530-0100) for 1 hour at room temperature. Primary antibodies were incubated overnight at 4°C at the following dilutions in blocking buffer: Ki-67 (Abcam, ab15580, 1:400), P62 (Abclonal, A19700, 1:1000), Cyp2f2 (SantaCruz, sc-374540, 1:200), Chop (Proteintech, 15204-1-AP, 1:200), Cyp1a2 (Santa Cruz, sc-53241, 1:50), p-eIF2A (CST, 3398S, 1:50), Atf4 (CST, 11815S, 1:200), and Btg2 (Proteintech, 22339-1-AP, 1:100). Secondary antibodies used were Goat anti-Rabbit IgG (H+L) Secondary Antibody, Biotin (Invitrogen, 65-6140, 1:1000) and peroxidase-AffiniPure Goat Anti-mouse IgG(H+L) (Jackson ImmunoResearch, 115-035-003, 1:1000). Avidin, NeutrAvidin™, Horseradish Peroxidase conjugate (Invitrogen, A2664, 1:2000) was used as the third antibody. After three washes with PBS, color development was achieved using the DAB Peroxidase Substrate Kit (ZSGB-BIO, ZLI-9018), followed by quenching in distilled water. Slides were counterstained with hematoxylin (BBI, A600701-0050), dehydrated to xylene, and images were captured using an Olympus BP80 microscope or Scan System SQS-1000 and analyzed with ImageJ.

Western blot analysis

Mouse liver tissue (30 mg) was homogenized in 1 mL RIPA buffer (10 mM Tris–HCl pH 8.0, 1 mM EDTA, 0.5 mM EGTA, 1% Triton X-100, 0.1% sodium deoxycholate, 0.1% SDS, 140 mM NaCl, 1 mM PMSF, and protease inhibitor cocktail). The lysate was mixed with 5× SDS loading buffer (final 1×), boiled at 98°C for 10 min, and centrifuged at 12,000 rpm for 3 min at room temperature. Supernatants were separated by 10% SDS-PAGE, transferred to PVDF membranes, and blocked with 5% nonfat milk in TBST (Tris-buffered saline with 0.1% Tween-20). Membranes were incubated with primary antibodies against Cyp2e1 (Proteintech, 19937-1-AP, 1:5000) and GAPDH (Proteintech, 60004-1-Ig, 1:5000), followed by HRP-conjugated secondary antibodies: goat anti-mouse (Jackson ImmunoResearch, 115-035-003, 1:2000) and goat anti-rabbit (Jackson ImmunoResearch, 111-035-003, 1:2000). Protein bands were detected using enhanced chemiluminescence (ECL) and a Minichemi Chemiluminescence Imaging System.

RNA extraction and quantitative real-time PCR

Ground mouse liver tissue was lysed in TRIzol reagent (Invitrogen, 15596018) and centrifuged at 12,000 rpm for 10 min at 4°C. The supernatant was transferred to a new tube, mixed with 200 μL chloroform, incubated at room temperature for 10 min, and centrifuged. The aqueous phase was collected, mixed with an equal volume of isopropanol, incubated at 4°C for 10 min, and centrifuged. The supernatant was discarded, and the RNA pellet was washed twice with 75% ethanol, air-dried for 10 min, and dissolved in RNase-free water. RNA concentration and purity were measured using a NanoDrop spectrophotometer. Reverse transcription was performed using a cDNA first-strand synthesis kit (Takara, 6210B). Quantitative PCR was carried out in triplicates using PowerUp™ SYBR™ Green Master Mix (Thermo Fisher, A25742) on a QuantStudio 1 Real-Time PCR system (Life Technologies) according to the manufacturer’s instructions. Primer sequences are provided in Table S1.

AAV production and purification

AAV8 was produced using 293T cells cultured in one or more 175 cm2 cell culture flasks. Cells were plated one day prior to transfection, reaching 50% confluence, which allowed them to reach 80-90% confluence the next day. For transfection of one 175 cm2 flask, a mixture of 11.56 μg pHelper, 7.24 μg pAAV2/8 (Addgene, 112864), and 5.9 μg transgene plasmids was prepared in 3 ml Opti-MEM medium in a tube. PEI solution (1 μg/μl in water, pH 7.0, powder from Polysciences, 23966-2) was added to the tube at a ratio of 3 μg PEI per 1 μg DNA. The solution was mixed and incubated for 15 minutes before adding to the cell culture. At 24 hours post-transfection, the medium was replaced with antibiotic-containing medium (P+/S+) without FBS. At 72 hours after the medium change, cells were scraped off the flask, and 3 ml of chloroform was added, followed by vortexing for 5 minutes. Then, 7.6 ml of 5M NaCl was added and vortexed for 10 seconds. The solution was centrifuged at 3000g for 5 minutes at 4°C. The supernatant was transferred to a new 50 ml tube, and 9.4 ml of 50% (w/v) PEG 8000 was added, vortexed for 10 seconds, and left for 1 hour at 4°C. Subsequently, the mixture was centrifuged at 3000g for 30 minutes at 4°C. The supernatant was decanted, and the tube was inverted for 10 minutes to dry. Next, 1.4 ml of 50 mM HEPES buffer (pH 8.0) was added to re-suspend the pellet, and the solution was vortexed for 5 minutes. Then, 3.5 μl of 1M MgCl2, 2.8 μl of 5 μg/μl DNase I, and 2.8 μl of 5 μg/μl RNase A were added, and the mixture was incubated for 20 minutes at 37°C. The virus was then extracted with chloroform, and the solvent was replaced with PBS. AAV was concentrated by centrifugation at 14000g for 5 minutes at 4°C using a 100kDa ultrafiltration tube (Beyotime, FUF058). Finally, the samples were aliquoted and stored at −80°C after titer measurement.

Library preparation for 10x Visium spatial transcriptomics

After administering APAP, mice were euthanized at 0h, 3h, 6h and 24h time points using carbon dioxide. Freshly harvested liver tissues were frozen using isopentane and liquid nitrogen, embedded in OCT, sectioned into 10 μm slices, and mounted on Visium tissue optimization or spatial gene expression slides. Tissues were fixed with methanol, stained with H&E, and imaged at 40× magnification. Tissue permeabilization was optimized for 6 minutes. Libraries were prepared according to the Visium spatial gene expression user guide, with cDNA amplification cycles adjusted per time point, and index PCR performed for 6 cycles. Libraries were pooled, loaded at 400 pM, and sequenced on a NovaSeq 6000 SP100 sequencer.

Processing of raw sequencing data

The raw data underwent processing using spaceranger (version 1.1.0) to perform several analytical tasks. These included detecting tissue boundaries, aligning reads to the reference genome, generating feature-spot matrices, conducting clustering to identify spatially distinct groups of cells, and analyzing gene expression patterns. Additionally, spaceranger facilitated the spatial placement of spots within the context of the slide image, providing a comprehensive spatial transcriptomic analysis of the tissue sample.

Basic analysis of spatial transcriptomics (ST) data

Earlier version: For the gene-spot matrices generated by Spaceranger, we utilized Scanpy (version 1.9.3) for analysis (Patro et al., 2017). Initially, routine statistical analyses were conducted, including calculating the number of detected UMIs (nUMI) and genes (nGene) in each spot. Basic quality control (QC) measures were then applied to the data. Specifically, spots with extremely low UMI or gene counts (min_counts=2000, min_cell=10) were excluded using scanpy.pp.filter_cells function.

Revised version: For the gene-spot matrices generated by Spaceranger, we utilized Scanpy (version 1.9.3) for analysis (Patro, Duggal et al. 2017). Initially, routine statistical analyses were conducted, including calculating the number of detected UMIs (nUMI) and genes (nGene) in each spot. Basic quality control (QC) measures were then applied to the data. Specifically, spots with extremely low UMI or gene counts (min_counts=2000, min_cell=10) were excluded using scanpy.pp.filter_cells function.

Mitochondrial and hemoglobin genes were filtered out to focus on relevant gene expression. Spot matrix was filtered out to keep only spots overlaying tissue sections. Following QC, data normalization was performed using the scanpy.pp.normalize_total function, and log-transformation was applied using the scanpy.pp.log1p function. Subsequently, we identified 2000 highly variable genes using the scanpy.pp.highly_variable_genes function based on their expression means and variances.

Principal component analysis (PCA) was then applied to reduce the dimensionality of the data, projecting the spots into a lower-dimensional space using the identified principal components (PCs). Using the corrected PC matrices, we conducted unsupervised clustering based on shared nearest neighbors (SNN) and visualized the data using UMAP (Uniform Manifold Approximation and Projection) for further exploration and analysis (Lim and Qiu 2023).

Each sample was partitioned into three clusters. To facilitate comparison, each cluster was annotated with a region label (PC, PP, or Mid), determined by integrating information from cluster-specific marker genes and H&E staining images. Furthermore, to compare gene expression profiles across clusters, we identified differentially expressed genes using fold change analysis and the Wilcoxon rank sum test. This approach allowed us to highlight genes showing significant expression differences among all or selected clusters, providing insights into biological variations across spatial regions in the tissue samples.

Analysis of differentially expressed genes (DEGs) and pathway enrichment

Earlier version: For differential gene analysis across different samples, we initially integrated all samples using the combat function to mitigate batch effects.(Stuart et al., 2019) Subsequently, differential expression analysis of variable genes was conducted using the FindMarkers function, employing the Wilcoxon rank sum test.(Ritchie et al., 2015) Genes showing significant changes were selected based on criteria including a log-fold change threshold of 0.1 and a minimum percentage of samples showing differential expression of 1%. The differentially expressed genes (DEGs) identified were enriched across various time points post-APAP injection, relative to the 0-hour time point.

Revised version: For differential gene analysis across different samples, we initially integrated all samples using the combat function to mitigate batch effects (Stuart, Butler et al. 2019). Subsequently, differential expression analysis of variable genes was conducted using the FindMarkers function, employing the Wilcoxon rank sum test (Ritchie, Phipson et al. 2015). Genes showing significant changes were selected based on criteria including a log-fold change threshold of 0.1 and a minimum percentage of samples showing differential expression of 1%. The differentially expressed genes (DEGs) identified were enriched across various time points post-APAP injection, relative to the 0-hour time point.

Earlier version: Pathway analysis was performed using the Kyoto Encyclopedia of Genes and Genomes (KEGG), while gene ontology (GO) analysis was carried out using the clusterProfiler package(Yu et al., 2012) to explore functional categories associated with these DEGs. The most significantly enriched pathways were identified using the enrichGO and enrichKEGG functions within the clusterProfiler package.(Yu et al., 2012) Furthermore, Gene Set Variation Analysis (GSVA) was employed to investigate the expression variations related to proliferation at different time points following APAP administration, utilizing the gsva function from the GSVA R package.(Hänzelmann et al., 2013) Pathway activities were visualized using the pheatmap R package for comprehensive analysis and interpretation. To functionally classify the DEGs identified in the F2A comparison, we performed Gene Ontology (GO) enrichment analysis using [name of tool/database, e.g., DAVID, clusterProfiler]. Each DEG was mapped to GO terms from the three ontologies: biological process, cellular component, and molecular function. When a single gene was annotated with multiple GO terms, we selected the most prominent term—defined as the one with the smallest adjusted p- value or highest enrichment score—for downstream interpretation. This approach ensured that the primary functional role of each DEG was captured without redundancy.

Revised version: Pathway analysis was performed using the Kyoto Encyclopedia of Genes and Genomes (KEGG), while gene ontology (GO) analysis was carried out using the clusterProfiler package (Yu, Wang et al. 2012) to explore functional categories associated with these DEGs. The most significantly enriched pathways were identified using the enrichGO and enrichKEGG functions within the clusterProfiler package (Yu, Wang et al. 2012). Furthermore, Gene Set Variation Analysis (GSVA) was employed to investigate the expression variations related to proliferation at different time points following APAP administration, utilizing the gsva function from the GSVA R package (Hänzelmann, Castelo et al. 2013). Pathway activities were visualized using the pheatmap R package for comprehensive analysis and interpretation. To functionally classify the DEGs identified in the F2A comparison, we performed Gene Ontology (GO) enrichment analysis using clusterProfiler. Each DEG was mapped to GO terms from the three ontologies: biological process, cellular component, and molecular function. When a single gene was annotated with multiple GO terms, we selected the most prominent term—defined as the one with the smallest adjusted p-value or highest enrichment score—for downstream interpretation. This approach ensured that the primary functional role of each DEG was captured without redundancy.

Gene regulatory network analysis

Earlier version: The Single-cell regulatory network inference and clustering (SCENIC) workflow comprises three main steps: coexpression analysis, target gene motif enrichment analysis, and regulon activity assessment. The analysis employed pySCENIC (version 0.12.1)(van de Sande et al., 2020) with default parameters, using the raw count matrix from all samples as input. Co-expression modules were initially identified, and the interaction strength between transcription factors (TFs) and their target genes was evaluated using GRNBoost.(Aibar et al., 2017) Each coexpressed module underwent motif enrichment analysis using RcisTarget with a rank threshold of 3000, to identify modules where the TF motif was significantly enriched among targets.(Aibar et al., 2017) Only modules meeting these criteria were retained, establishing TFs along with their potential direct targets as regulons. The activity of each regulon within each cell was assessed using AUCell. For visualization, average regulon activity scores (AUC) were calculated for each cluster.(Aibar et al., 2017) A rank plot of regulons was generated using ggplot2. Additionally, the gene expression network related to the target genes of the TFs of interest was visualized using the nx.draw function.

Revised version: The Single-cell regulatory network inference and clustering (SCENIC) workflow comprises three main steps: coexpression analysis, target gene motif enrichment analysis, and regulon activity assessment. The analysis employed pySCENIC (version 0.12.1)(van de Sande, Flerin et al. 2020) with default parameters, using the raw count matrix from all samples as input. Co-expression modules were initially identified, and the interaction strength between transcription factors (TFs) and their target genes was evaluated using GRNBoost (Aibar, González-Blas et al. 2017). Each coexpressed module underwent motif enrichment analysis using RcisTarget with a rank threshold of 3000, to identify modules where the TF motif was significantly enriched among targets (Aibar, González-Blas et al. 2017). Only modules meeting these criteria were retained, establishing TFs along with their potential direct targets as regulons. The activity of each regulon within each cell was assessed using AUCell. For visualization, average regulon activity scores (AUC) were calculated for each cluster (Aibar, González-Blas et al. 2017). A rank plot of regulons was generated using ggplot2. Additionally, the gene expression network related to the target genes of the TFs of interest was visualized using the nx.draw function.

Analysis of snRNA-seq Data (GSE223561)

To annotate cell clusters in the single-nucleus RNA-seq (snRNA-seq) data, we utilized a curated list of known marker genes corresponding to liver cell lineages, as reported in Multimodal Decoding of Human Liver Regeneration. Using the AddModuleScore function in Seurat, we calculated signature scores to assign cell lineage identities. Clusters identified as primarily composed of cycling cells were further reclustered to segregate them into their constituent lineages. This process was repeated iteratively for each identified lineage. A cleansing step was included to remove clusters enriched for nuclei identified as doublets or overexpressing marker genes from multiple lineages. Hepatocytes were subsequently extracted for further reclustering, using resolution (res) = 0.2 and 80 principal components (npc). To annotate the hepatocyte clusters, we applied a curated list of known zonation-specific marker genes for hepatocytes.

Library preparation for Cut&Run samples

Cut&Run was conducted on mouse liver tissue samples from mice injected with APAP for 6 hours using the Cut&Run assay kit (CST, 86652) following the manufacturer’s protocol. Initially, 1 mg of fresh tissue was weighed for each antibody/MNase reaction, with an additional 1 mg for the input sample. Reactions included positive and negative controls. Tissues were fixed with 1 ml of 0.1% formaldehyde for 10 minutes and quenched with glycine for 5 minutes. After washing with PBS, tissues were resuspended in 1 ml of 1X Wash Buffer (+ spermidine + protease inhibitor cocktail) and transferred to a Dounce homogenizer. Tissue was homogenized into a single-cell suspension with 10-15 strokes until no tissue chunks remained.

For the Cut&Run assay to identify Chop binding sites, samples were bound to concanavalin A beads and incubated overnight at 4°C with the following primary antibodies: 0.5 μg CHOP (CST, 2895S, 1:100), 2 μl positive control Tri-Methyl-Histone H3 (Lys4) (CST, 9751), and 5 μl negative control Rabbit (DA1E) mAb IgG XP® Isotype Control (CST, 66362). The pAG MNase enzyme was activated by adding cold calcium chloride and incubating at 4°C for 30 minutes. The reaction was terminated with 1X stop buffer at 37°C for 10 minutes without shaking to release DNA fragments into the solution.

To reverse crosslinks in fixed tissue samples, samples were brought to room temperature and treated with 10% SDS Solution (BBI, B548118-0100) to a final concentration of 0.1%, along with proteinase K (20 mg/ml). DNA Extraction Buffer (+ Proteinase K + RNAse A) was added to the input sample. Cells were lysed and chromatin fragmented by sonication using a Covaris ME220 contact ultrasound apparatus on ice, with optimal conditions generating chromatin fragments ranging from 100 to 600 bp. DNA was purified using phenol/chloroform extraction followed by ethanol precipitation, and DNA quantification was performed by qPCR. For Cut&Run, purified DNA was used to prepare sequencing libraries with the NEB Next Ultra II DNA Library Prep Kit for Illumina (7645S and 7335L) and sequenced on an Illumina Nova-Seq 6000 sequencer to obtain 150 bp paired-end reads.

Cut&Run data analysis

Earlier version: For visualization purposes, BW files were generated using Integrative Genomics Viewer (IGV) software (Robinson et al., 2011), and bigwig files were created using Deeptools (version 3.5.1) with the ‘bamCoverage’ module (bamCoverage --binSize 100) (Ramírez et al., 2016). For each biological replicate and its corresponding IgG control, peaks were called using macs2 (version 2.2.7) with the command macs2 callpeak -f BAMPE --qvalue 0.1 --keep-dup (Zhang et al., 2008). To annotate the identified peaks, ChIPseeker (version 1.26.2) in R was utilized (Yu et al., 2015).

Revised version: For visualization purposes, BW files were generated using Integrative Genomics Viewer (IGV) software (Robinson, Thorvaldsdóttir et al. 2011), and bigwig files were created using Deeptools (version 3.5.1) with the ‘bamCoverage’ module (bamCoverage --binSize 100) (Ramírez, Ryan et al. 2016). For each biological replicate and its corresponding IgG control, peaks were called using macs2 (version 2.2.7) with the command macs2 callpeak -f BAMPE --qvalue 0.1 --keep-dup (Zhang, Liu et al. 2008). To annotate the identified peaks, ChIPseeker (version 1.26.2) in R was utilized (Yu, Wang et al. 2015).

Quantification of zonal distribution

Earlier version: The zonation of protein (such as Ki67, Chop, Atf4) -positive nuclei was quantified using the following method: the position index (P.I.) was calculated based on distances to the nearest CV (x), portal vein (PV) (y), and the distance between CV and PV (z), utilizing the law of cosines. The formula employed was P.I. = (x^2 + z^2 - y^2) / (2z^2). This approach aligns with the methodology described by Lin et al (Lin et al., 2018).

Revised version: The zonation of protein (such as Ki67, Chop, Atf4) -positive nuclei was quantified using the following method: the position index (P.I.) was calculated based on distances to the nearest CV (x), portal vein (PV) (y), and the distance between CV and PV (z), utilizing the law of cosines. The formula employed was P.I. = (x^2 + z^2 - y^2) / (2z^2). This approach aligns with the methodology described by Lin et al (Lin, Nascimento et al. 2018).

Statistical analysis

All experimental data are presented as the mean ± SD. Statistical analyses were performed using GraphPad Prism (v8.0.2), with appropriate tests selected based on experimental design: unpaired two-tailed Student’s t-tests for comparisons between two groups, and one-way ANOVA for comparisons involving three or more groups. The number of animals (‘‘n’’) used in each experiment is indicated in the Figures and corresponding legends. For quantification of liver sections, three to five random pericentral and periportal fields of each liver sample, unless specified otherwise, were imaged and quantified using Image J.